Neural Correlates and Reinstatement of Recent and Remote Memory: A Comparison Between Children and Young Adults

Author response:

The following is the authors’ response to the original reviews.

Reviewer #1 (Public Reviews):

Summary:

This paper by Schommartz and colleagues investigates the neural basis of memory reinstatement as a function of both how recently the memory was formed (recent, remote) and its development (children, young adults). The core question is whether memory consolidation processes as well as the specificity of memory reinstatement differ with development. A number of brain regions showed a greater activation difference for recent vs. remote memories at the long versus shorter delay specifically in adults (cerebellum, parahippocampal gyrus, LOC). A different set showed decreases in the same comparison, but only in children (precuneus, RSC). The authors also used neural pattern similarity analysis to characterize reinstatement, though I have substantive concerns about how this analysis was performed and as such will not summarize the results. Broadly, the behavioural and univariate findings are consistent with the idea that memory consolidation differs between children and adults in important ways, and takes a step towards characterizing how.

Strengths:

The topic and goals of this paper are very interesting. As the authors note, there is little work on memory consolidation over development, and as such this will be an important data point in helping us begin to understand these important differences. The sample size is great, particularly given this is an onerous, multi-day experiment; the authors are to be commended for that. The task design is also generally well controlled, for example as the authors include new recently learned pairs during each session.

Weaknesses:

As noted above, the pattern similarity analysis for both item and category-level reinstatement was performed in a way that is not interpretable given concerns about temporal autocorrelation within the scanning run. Below, I focus my review on this analytic issue, though I also outline additional concerns.

We thank the reviewer for both the positive and critical appraisal of our paper.

(1) The pattern similarity analyses were not done correctly, rendering the results uninterpretable (assuming my understanding of the authors' approach is correct).

a. First, the scene-specific reinstatement index: The authors have correlated a neural pattern during a fixation cross (delay period) with a neural pattern associated with viewing a scene as their measure of reinstatement. The main issue with this is that these events always occurred back-to-back in time. As such, the two patterns will be similar due simply to the temporal autocorrelation in the BOLD signal. Because of the issues with temporal autocorrelation within the scanning run, it is always recommended to perform such correlations only across different runs. In this case, the authors always correlated patterns extracted from the same run, which moreover have temporal lags that are perfectly confounded with their comparison of interest (i.e., from Fig 4A, the "scene-specific" comparisons will always be back-to-back, having a very short temporal lag; "set-based" comparisons will be dispersed across the run, and therefore have a much higher lag). The authors' within-run correlation approach also yields correlation values that are extremely high - much higher than would be expected if this analysis was done appropriately. The way to fix this would be to restrict the analysis to only cross-run comparisons, but I don't believe this is possible unfortunately given the authors' design; I believe the target (presumably reinstated) scene only appears once during scanning, so there is no separate neural pattern during the presentation of this picture that they can use. For these reasons, any evidence for "significant scene-specific reinstatement" and the like is completely uninterpretable and would need to be removed from the paper.

We thank the reviewer for this important input. We acknowledge that our study design leads to temporal autocorrelation in the BOLD signal when calculating RSA between fixation and scene time windows. We also recognize that we cannot interpret the significance of scene-specific reinstatement compared to zero and have accordingly removed this information. Nevertheless, our primary objective was to investigate changes in scene-specific reinstatement in relation to the different time delays of retrieval. Given that the retrieval procedure is the same over time and presumably similarly influenced by temporal autocorrelations, we argue that our results must be attributed to the relative differences in reinstatement across recent and remote trials. Bearing this in mind, we argue that our results can be interpreted in terms of delay-related changes in reinstatement. This information is discussed in pp. 21, 40 of the manuscript.

We agree with the reviewer that cross-run comparisons would be extremely interesting. This could be achieved by introducing the same items repeatedly across different runs, which was not possible in our current setup since we were interested in single exposure retrieval and practical time restriction in scanning children. We have introduced this idea in Limitations and Discussion sections (pp. 40, 44) of the manuscript to inform future studies.

Finally, thanks to the reviewer’s comment, we identified a bug in the final steps of our RSA calculation. Fischer’s z-transformation was incorrectly applied to r-1 values, resulting in abnormally high values. We apologize for this error. We have revised the scripts and rectified the bug by correctly applying Fischer’s z-transformation to the r similarity values. We also adjusted the methods description figure accordingly (Figure 5, p. 22). This adjustment led to slightly altered reinstatement indices. Nevertheless, the overall pattern of delay-related attenuation in the scene-specific reinstatement index, observed in both children and adults, remains consistent. Similarly, we observed gist-like reinstatement uniquely in children.

b. From a theoretical standpoint, I believe the way this analysis was performed considering the fixation and the immediately following scene also means that the differences between recent and remote could have to do with either the reactivation (processes happening during the fixation, presumably) or differences in the processing of the stimulus itself (happening during the scene presentation). For example, people might be more engaged with the more novel scenes (recent) and therefore process those scenes more; such a difference would be interpreted in this analysis as having to do with reinstatement, but in fact could be just related to the differential scene processing/recognition, etc.

Thank you for your insightful comments. We acknowledge the theoretical concerns raised about distinguishing between the effects of reactivation processes occurring during fixation and differential processing of the stimulus itself during scene presentation. Specifically, the notion that engagement levels with recent scenes could result in enhanced processing, which might be misattributed to memory reinstatement mechanisms.

We argue, however, that during scene presentation, scenes are processed more “memory-wise” rather than “perception-wise”, since both recent and remote memories are well-learned, as we included only correctly recalled memories in the analysis.

We concur that scene presentations entail perceptual processing; however, such processing would be consistent across all items, given that they were presented with the same repeated learning procedure, rendering them equally familiar to participants. In addition, we would argue that distinct activation patterns elicited during varying delays are more likely attributable to memory-related processing, since participants actively engaged in a memory-based decision-making task during these intervals. We have incorporated this rationale into the discussion section of our manuscript (p. 40).

With this in mind, we hypothesized that in case of “memory-wise” processing, the neural engagement during the scene time window should be higher for remote compared to recent items, and this increases with passing time as more control and effort should be exhibited during retrieval due to reorganized and distributed nature of memories. If the scenes are processed more “perception-wise”, we would expect higher neural engagement during the retrieval of recent compared to remote items. Our exploratory analysis (detailed overview in supplementary materials, Figure S3, Table S9) revealed a higher neural activation for remote compared to recent items in medial temporal, prefrontal, occipital and cerebellar brain regions, supporting the notion of “memory-wise” processes during scene time window. However, this exploratory analysis cannot provide a direct solution to the reviewer’s concern as our paradigm per se cannot arbitrate between “memory-wise” and “perception-wise” nature of retrieval. We added the point to the discussion (see p. 40).

c. For the category-based neural reinstatement:

(1) This suffers from the same issue of correlations being performed within the run. Again, to correct this the authors would need to restrict comparisons to only across runs (i.e., patterns from run 1 correlated with patterns for run 2 and so on). With this restriction, it may or may not be possible to perform this analysis, depending upon how the same-category scenes are distributed across runs. However, there are other issues with this analysis, as well.

(2) This analysis uses a different approach of comparing fixations to one another, rather than fixations to scenes. The authors do not motivate the reason for this switch. Please provide reasoning as to why fixation-fixation is more appropriate than fixation-scene similarity for category-level reinstatement, particularly given the opposite was used for item-level reinstatement. Even if the analyses were done properly, it would remain hard to compare them given this difference in approach.

(3) I believe the fixation cross with itself is included in the "within category" score Is this not a single neural pattern correlated with itself, which will yield maximal similarity (pearson r=1) or minimal dissimilarity (1-pearson r=0)? Including these comparisons in the averages for the within-category score will inflate the difference between the "within-category" and "between-category" comparisons. These (e.g., forest1-forest1) should not be included in the within-category comparisons considered; rather, they should be excluded, so the fixations are always different but sometimes the comparisons are two retrievals of the same scene type (forest1-forest2), and other times different scene types (forest1-field1)

(4) It is troubling that the results from the category reinstatement metric do not seem to conceptually align with past work; for example, a lot of work has shown category-level reinstatement in adults. Here the authors do not show any category-level reinstatement in adults (yet they do in children), which generally seems extremely unexpected given past work and I would guess has to do with the operationalization of the metric.

Thank you for this important input regarding category-based reinstatement.

(1) The distribution of within-category items across runs was approximately similar and balanced. Additionally, within runs, they were presented randomly without close temporal proximity. Based on this arrangement, we believe that the issue of close temporal autocorrelation, as pointed out by the reviewer in the context of scene-specific reinstatement, may not apply to the same extent here. Again, our focus is not on the absolute level of category-based reinstatement, but the relative difference across conditions (recent vs. remote short delay vs. remote long delay) which are equally impacted by the autocorrelations.

(2) We apologize for not motivating this analysis further. Whereas the scene-reinstatement index (i.e., fixation to scene correlation) gives us a measure of the pre-activation of a concrete scene (e.g., a yellow forest in autumn), the gist-like reinstatement gives us a measure of the pre-activation of a whole category of scenes (e.g., forests). Critically, our window of interest is the fixation period for both sets of analysis (in the absence of any significant visual input). The scene-specific reinstatement uses the scene window as a neural template against which the fixation period can be compared, while the gist-like reinstatement compares similarity of reactivation pattern for trials from the same category but differ in the exact memory content. The reinstatement of more generic, gist-like memory (e.g., forest) across multiple trials should yield more similar neural activation patterns. Significant gist-like reinstatement would suggest that neural patterns for scenes within the same category are more generic, as indicated by higher similarity among them. On the other hand, a more detailed reinstatement of specific types of forests (e.g., a yellow forest in autumn, green pine trees, a bare-leaved forest in spring, etc.) that differ in various dimensions could result in neural activation patterns that are as dissimilar as those seen in the reinstatement of scenes from entirely different categories. Through this methodology, we could distinguish between more generic, gist-like reinstatement and more specific, detailed reinstatement. This is now clarified in the manuscript, see p. 25.

(3) We apologize for the confusion caused by the figure and analysis description. In our analysis, we indeed excluded the correlation of the fixation cross with itself. Consequently, the diagonal in the figure should be blank to indicate this. This is now revised in the manuscript (Figure 7B and in Methods).

(4) We appreciate your concern and recognize that the terminology we used might not align perfectly with the conventional understanding of category-based reinstatement. Typically, category-level neural representations (as discussed in Polyn et al., 2005; Jafarpour et al., 2014; among others) are investigated to identify specific brain areas associated with encoding/perception of scenes or faces. Our aim, however, was to explore the mnemonic reinstatement of highly detailed scenes that were elaborately encoded, with the hypothesis that substantial representational transformations would occur over time and vary with age. This hypothesis is based on the memory literature, including the Fuzzy-Trace Theory, the Contextual Binding Theory, and the Trace Transformation Theory (Brainerd & Reyna, 1998; Yonelinas, 2019; Moscovitch & Gilboa, 2023). Therefore, we renamed 'category-based' reinstatement to 'gist-like' reinstatement, which clarifies our concept and better aligns it with existing literature.

We anticipated that young adults, having the ability to retain detailed narratives post-encoding, would demonstrate a reinstatement of scenes with distinct details, making these scenes dissimilar from each other (see similar findings in Sommer et al., 2021). In contrast, given the anticipated lesser strategic elaboration during learning in children, we hypothesized that they would demonstrate a shallower, more gist-like reinstatement (for instance, children recalling a forest or a field in a general sense without specific details or vivid imagery). This could result in higher category-based similarity, as children might reinstate a more generic forest concept.

We did not gather additional data on the verbal quality of reinstatement due to the limited scanning time available for children, so these assumptions remain unverified. However, anecdotal observations post-retrieval indicated that adults often reported very vivid scenes associated with clear narrative recall. In contrast, children frequently described more vague memories (e.g., “I know it was a forest”) without specific details. Future studies should include measures to assess the quality of reinstatement, potentially outside the scanning environment.

(2) I did not see any compelling statistical evidence for the claim of less robust consolidation in children.

Specifically in terms of the behavioral results of retention of the remote items at 1 vs 14 days, shown in Figure 2B, the authors conclude that memory consolidation is less robust in children (line 246). Yet they do not report statistical evidence for this point, as there was no interaction of this effect with the age group. Children had worse memory than adults overall (in terms of a main effect - i.e. across recent and remote items). If it were consolidation-specific, one would expect that the age differences are bigger for the remote items, and perhaps even most exaggerated for the 14-day-old memories. Yet this does not appear to be the case based on the data the authors report. Therefore, the behavioral differences in retention do not seem to be consolidation specific, and therefore might have more to do with differences in encoding fidelity or retrieval processes more generally across the groups. This should be considered when interpreting the findings.

Thank you for highlighting this important issue. We acknowledge that our initial description and depiction of our behavioral findings may not have effectively conveyed the main message about memory consolidation. Therefore, we have revised the behavioral results section (see pp. 12-14) to communicate our message more clearly.

As detailed in the methods section, we reported retention rates only for those items that were correctly (100%) learned on day 0, day 1, and day 14. This approach meant that different participants had varying numbers of items learned correctly. However, this strategy allowed us to address our primary question: whether memory consolidation, based on all items initially encoded successfully, is comparably robust between groups.

To illustrate the change in retention rate slopes over time for recently learned items (i.e., immediately 30 minutes after learning), short delay remote, and long delay remote items, relative to the initially correctly learned items more clearly and straightforward, we conducted the following analysis: after observing no differences between sessions in both age groups for recent items on days 1 and 14, we combined the recent items. This approach enabled us to investigate how the slope of memory retention for initially correctly learned items (with a baseline of 100%) changes over time. We observed a significant interaction between item type (recent, short delay remote, long delay remote) and group (F(3,250) = 17.35, p < .001, w2 = .16). The follow up of this interaction revealed significantly less robust memory consolidation across all delay times in children compared to young adults. This information is added in the manuscript in pp. 12-14. We have also updated the figures, incorporating the baseline of 100% correct performance.

(3) Please clarify which analyses were restricted to correct retrievals only. The univariate analyses states that correct and incorrect trials were modelled separately but does not say which were considered in the main contrast (I assume correct only?). The item specific reinstatement analysis states that only correct trials were considered, but the category-level reinstatement analysis does not say. Please include this detail.

Thank you for bringing this to our attention. We indeed limited our analysis – including univariate, specific reinstatement, and gist-like analyses – to only correctly remembered items. This decision was made because our goal was to observe delay-related changes in the neural correlates of correct memories, which are potentially stronger. We have incorporated this information into the manuscript.

(4) To what extent could performance differences be impacting the differences observed across age groups? I think (see prior comment) that the analyses were probably limited to correct trials, which is helpful, but still yields pretty big differences across groups in terms of the amount of data going into each analysis. In general, children showed more attenuated neural effects (e.g., recent/remote or session effects); could this be explained by their weaker memory? Specifically, if only correct trials are considered that means that fewer trials would be going into the analysis for kids, especially for the 14-day remote memories, and perhaps pushing the remove > recent difference for this condition towards 0. The authors might be able to address this analytically; for example, does the remote > recent difference in the univariate data at day 14 correlate with day 14 memory?

Thank you for pointing this out. Indeed, there was a significant relationship between remote > recent difference in the univariate data and memory performance at day 14 across both age group (see Figure 4C-D). The performance of all participants including children was above chance level for remote trial on day 14. In addition, although number of remote trials was lower in children (18 trials on average) in comparison to adults (22 trials on average), we believe that the number of remote trials was not too low or different across groups for the contrast.

(5) Some of the univariate results reporting is a bit strange, as they are relying upon differences between retrieval of 1- vs. 14-day memories in terms of the recent vs. report difference, and yet don't report whether the regions are differently active for recent and remote retrieval. For example, in Figure 3A, neither anterior nor posterior hippocampus seem to be differentially active for recent vs. remote memories for either age group (i.e., all data is around 0). This difference from zero or lack thereof seems important to the message - is that correct? If so, can the authors incorporate descriptions of these findings?

Thank you for this valuable input. When examining recent and remote retrieval separately, indeed both the anterior and posterior regions of the hippocampus exhibited significant activation from zero in adults (all p < .0003FDRcorr) and children (all p < .014FDRcorr, except for recent posterior hippocampus) during all delays. We include this information in the manuscript (see p. 17) and add it to the supplementary materials (Figure S2, Table S7).

(6) Please provide more details about the choices available for locations in the 3AFC task. (1) Were they different each time, or always the same? If they are always the same, could this be a motor or stimulus/response learning task? (2) Do the options in the 3AFC always come from the same area - in which case the participant is given a clue as to the gist of the location/memory? Or are they sometimes randomly scattered across the image (in which case gist memory, like at a delay, would be sufficient for picking the right option)? Please clarify these points and discuss the logic/impact of these choices on the interpretation of the results.Response: Thank you for pointing this out. During learning and retrieval, we employed the 3AFC (Three-Alternative Forced Choice) task.

The choices for locations varied across scenes while remained the same across time within individuals. There were 18 different key locations for the objects, distributed across the stimulus set. This means the locations of the objects were quite heterogeneous and differed between objects. The location of the object within the task was presented once during encoding and remained consistent throughout learning. Given the location heterogeneity, we believe our task cannot be reduced to a mere “stimulus/response learning task” but is more accurately described as an object-location associations task.

Similar to the previous description, the options for the 3AFC task did not originate from the same area, as there were 18 different areas in total. The three choice options were distributed equally: so sometimes the “correct” answer was the left option, sometimes in the middle option, or sometimes the right option. Therefore, we believe that the 3AFC task did not provide clues to the location but required detailed and precise memory of the location. Moreover, the options were not randomly scattered but rather presented close together in the scene, demanding a high level of differentiation between choices.

Taking all the above into consideration, we assert that precise object-location associative memory is necessary for a correct answer. We have added this information to the manuscript (p. 9).

(7) Often p values are provided but test statistics, effect sizes, etc. are not - please include this information. It is at times hard to tell whether the authors are reporting main effects, interactions, pairwise comparisons, etc.

Thank you for bringing this to our attention. We realize that including this information in the Tables may not be the most straightforward approach. Therefore, we have incorporated the test statistics, effect sizes, and related details into the text of the results section for clarity.

(8) There are not enough methodological details in the main paper to make sense of the results. For example, it is not clear from reading the text that there are new object-location pairs learned each day.

Thank you for pointing this out. We have added this information to the main manuscript. Additionally, we have emphasized this information in the text referring to Figure 1B.

(9) The retrieval task does not seem to require retrieval of the scene itself, and as such it would be helpful for the authors to both explain their reasoning for this task to measure reinstatement. Strictly speaking, participants could just remember the location of the object on the screen. Was it verified that children and adults were recalling the actual scene rather than just the location (e.g. via self-report)? It's possible that there may be developmental differences in the tendency to reinstate the scene depending on e.g., their strategy.

Thank you for highlighting this important point. Indeed, the retrieval task included explicit instructions for participants to recall and visualize the scene associated with the object presented during the fixation time window. Participants were also instructed to recollect the location of the object within the scene. Since the location was contextually bound to the scene and each object had a unique location in each scene, the location of the object was always embedded in the specific scene context. We have added this information to both the Methods and Results sections.

From the self-reports of the participants (which unfortunately were not systematically collected on all occasions), they indicated that when they could recall the scene and the location due to the memory of stories created during strategic encoding, it aided their memory for the scene and location immensely. We also concur with your observation that children and young adults may differ in their ability to reinstate scenes, depending on the success of their employed recall strategies. This task was conducted with an awareness of potential developmental differences in the ability to form complex contextual memories. Our elaborative learning procedure was designed to minimize these differences. It is important to note though we did not expect children to achieve performance levels fully comparable to adults. There may indeed be developmental differences in reinstatement, such as due to differences in knowledge availability and accessibility (Brod, Werkle-Bergner, & Shing, 2013). We think that these differences may underlie our findings of neural reinstatement. This is now discussed in p. 34-35, 39-43 of the manuscript.

(10) In general I found the Introduction a bit difficult to follow. Below are a few specific questions I had.

a. At points findings are presented but the broader picture or take-home point is not expressed directly. For example, lines 112-127, these findings can all be conceptualized within many theories of consolidation, and yet those overarching frameworks are not directly discussed (e.g., that memory traces go from being more reliant on the hippocampus to more on the neocortex). Making these connections directly would likely be helpful for many readers.

Thank you for bringing this to our attention. We have incorporated a summary of the general frameworks of memory consolidation into the introduction. This addition outlines how our summarized findings, particularly those related to memory consolidation for repeatedly learned information, align with these frameworks (see lines 126-138, 146-150).

b. Lines 143-153 - The comparison of the Tompary & Davachi (2017) paper with the Oedekoven et al. (2017) reads like the two analyses are directly comparable, but the authors were looking at different things. The Tompary paper is looking at organization (not reinstatement); while the Oedekoven et al. paper is measuring reinstatement (not organization). The authors should clarify how to reconcile these findings.

Thank you for highlighting this aspect. We have revised how we present the results from Tompary & Davachi (2017). This study examined memory reorganization for memories both with and without overlapping features, and it observed higher neural similarity for memories with overlapping features over time. The authors also explored item-specific reinstatement for recent and remote memories by assessing encoding-retrieval similarity. Since Oedekoven et al. (2017) utilized a similar approach, their results are comparable in terms of reinstatement. We have updated and expanded our manuscript to clarify the parallels between these studies (see lines 157-162).

c. Line 195-6: I was confused by the prediction of "stable involvement of HC over time" given the work reviewed in the Introduction that HC contribution to memory tends to decrease with consolidation. Please clarify or rephrase.

Drawing on the Contextual Binding Theory (Yonelinas et al., 2019), as well as the Multiple Trace Theory (Nadel et al., 2000) and supported for instance by evidence from Sekeres et al. (2018), we hypothesized that detailed contextual memories formed through repeated and strategic learning would strengthen the specificity of these memories, resulting in consistent hippocampal involvement for successfully recalled contextualized detailed memories. We have included additional explanatory information in the manuscript to clarify this hypothesis (see lines 217-219).

d. Lines 200-202: I was a bit confused about this prediction. Firstly, please clarify whether immediate reinstatement has been characterized in this way for kids versus adults. Secondly, don't adults retain gist more over long delays (with specific information getting lost), at least behaviourally? This prediction seems to go against that; please clarify.

Thank you for raising this important point. Indeed, there are no prior studies that examined memory reinstatement over extended durations in children. The primary existing evidence suggests that neural specificity or patterns of neural representations in children can be robustly observed, while neural selectivity or univariate activation in response to the same stimuli tends to mature later (i.e., Fandakova et al., 2019). Bearing this in mind and recognizing that such neural patterns can be observed in both children and adults, we hypothesized that adults may form stronger detailed contextual memories compared to children. By employing strategies such as creating stories, adults might more easily recall scenes without the need to resort to forming generic or gist-like memories (for example, 'a red fox was near the second left pine tree in a spring green forest'). This assumption aligns with the Fuzzy Trace Theory (Reyna & Brainerd, 1995), which posits that verbatim memories can be created without the extraction of a gist.

Conversely, we hypothesized that children, due to the ongoing maturation of associative and strategic memory components (as discussed in Shing et al., 2008 and 2010), which are dependent respectively on the hippocampus (HC) and the prefrontal cortex (PFC), would be less adept at creating, retaining, and extracting stories to aid their retrieval process. This could result in them remembering more generic integrated information, like the relationship between a fox and some generic image of a forest. We have added explanatory information to the manuscript to elucidate these points (see lines 225-230).

Reviewer #1 (Recommendations For The Authors):

(1) For Figure 3, I would highly recommend changing the aesthetics for the univariate data - at least on my screen they appear to be open boxes with solid vs. dashed lines, and as such look identical to the recent vs. remove distinction in Figure 2B. It also doesn't match the legend for me, which shows the age groups having purple vs. yellow coloring.

Thank you for this observation. We have adjusted Figure 2 (now Figure 3) (please refer to p. 14) accordingly, now utilizing purple and yellow colors to distinguish between the age groups.

(2) Lines 329-330, it is not true that "all" indices were significant from zero but this is only apparent if you read the next sentence. Please rephrase to clarify. e.g., "All ... indices with a few exceptions ... were significantly..."?

Based on the above suggestions and considering our primary focus on time-related changes in scene-specific reinstatement, we will refrain from further interpreting the relative expression of individual scene-specific indices against 0. Consequently, we have removed this information from our analysis.

(3) It is challenging to interpret some of the significance markers, such as those in Figure 3. For example what effects are being denoted by the asterisks and bars above vs. below the data on panel D? Please clarify and/or note in the legend.

We have included a note in the legend to clarify the meaning of all significance markers. In addition, we decided to state any significant main and interaction effects in the figure rather that to use significance markers.

(4) For Figures 2 and 3, only the meaning of error bars is described in the caption. It is not explained in the caption what the boxes, lines, and points denote. Please clarify.

Thank you for highlighting this. We have added explanations to the figure's annotation for clarity. Please note, that considering other review’s suggestions figure plots may have been adjusted or changed, resulting in adjustment of the explanations in the figure annotation.

(5) How were recent and remote interspersed relative to one another? The text says that each run had 10 recent and 10 remote pairs, presented in a "pseudo-random order" - not clear what that (pseudo) means in this case. Please clarify.

Thank you for raising this point. We provide this information in the Methods section “Materials and Procedure”: 'The jitters and the order of presentation for recent and remote items were determined using OptimizeXGUI (Spunt, 2016), following an exponential distribution (Dale, 1999). Ten unique recently learned pairs (from the same testing day) and ten unique remotely learned items (from Day 0) were distributed within each run (in total three runs) in the order as suggested by the software as the most optimal. There were three runs with unique sets of stimuli each resulting in thirty unique recent and thirty unique remote stimuli overall.'

(6) Figure 1A, second to last screen on the learning cycles row - what would be presented to participants here, one of these three emojis? What does the sleepy face represent? I see some of these points were mentioned in the methods, but additional clarification in the caption would be helpful.

Thank you for highlighting this. We have included this information in the figure caption. Specifically, the sleepy face symbol in the figure denotes a 'missed response'.

(7) Not clear how the jittered fixation time between object presentation and scene test is dealt with in representational similarity analyses.

Thank you for pointing this out. Beta estimates were obtained from a Least Square Separate (LSS) regression model. Each event was modeled with their respective onset and duration and, as such, one beta value was estimated per event (with the lags between events differing from trial to trial). We have edited the corresponding section (see p. 53).

(8) It was a little bit strange to have used anterior vs posterior HPC ROIs separately in univariate analysis but then combined them for multivariate. There are many empirical and theoretical motivations for looking at item-specific and category reinstatement in anterior and posterior HPC separately, so I was surprised not to see this. Please explain this reasoning.

Thank you for pointing this out. We agree with the reviewer and included the anterior and posterior HC ROIs into the multivariate analysis. Please see the revised results section (pp. 13-15).

(9) The term "neural specificity" is introduced (line 164) without explanation; please clarify.

Thank you for bringing this to our attention. The term ‘neural specificity’ refers to the neural representational distinctiveness of information. In other words, ‘neural specificity,’ as defined by Fandakova et al. (2019), refers to the distinctiveness of neural representations in the regions that process that sensory input. We decided, however to refrain from using this term and instead to use neural representational distinctiveness, which is more self-explaining and was also introduced in the manuscript.

(10) Age range is specified as 5-7 years initially (line 187) and then 6-7 years (line 188).

We have corrected the age range in line 188 to '5 to 7 years.'

Reviewer #2 (Public Reviews):

Schommartz et al. present a manuscript characterizing neural signatures of reinstatement during cued retrieval of middle-aged children compared to adults. The authors utilize a paradigm where participants learn the spatial location of semantically related item-scene memoranda which they retrieve after short or long delays. The paradigm is especially strong as the authors include novel memoranda at each delayed time point to make comparisons across new and old learning. In brief, the authors find that children show more forgetting than adults, and adults show greater engagement of cortical networks after longer delays as well as stronger item-specific reinstatement. Interestingly, children show more category-based reinstatement, however, evidence supports that this marker may be maladaptive for retrieving episodic details. The question is extremely timely both given the boom in neurocognitive research on the neural development of memory, and the dearth of research on consolidation in this age group. Also, the results provide novel insights into why consolidation processes may be disrupted in children. Despite these strengths, there are quite a few important design and analytical choices that derail my enthusiasm for the paper. If the authors could address these concerns, this manuscript would provide a solid foundation to better understand memory consolidation in children.

We thank the reviewer for both the positive and critical appraisal of our paper.

Reviewer #2 (Recommendations For The Authors):

(1) My greatest concern is the difference in memory accuracy that emerges as soon as immediate learning, which undermines the interpretation of any consolidation-related differences. This concern is two-fold. The authors utilize an adaptive learning approach in which participants learn to criteria or stop after 4 repetitions. This type of approach leads to children seeing the stimuli more often during learning compared to adults, which on its own could have consequences for consolidation-related neural markers. Specifically, within adults theoretical and empirical work this shows that repeating information can actually lead to more gist-like representations, which is the exact profile the children are showing. While there could be a strength to this approach because it allows for equivocal memory, the decision to stop repetitions before criteria means that memory performance is significantly lower in the children, which again could have consequences to consolidation-related neural markers. First, the authors do not show any of the learning-related data which would be critical to assess the impact of this design choice. Second, there are likely differences in memory strength at the delay, making it extremely difficult to determine if the neural markers reflect development, worse memory strength, or both. This issue is compounded by the use of a 3-AFC paradigm, wherein "correct responses" included in the analysis could contain a significant amount of guessing responses. I think a partial solution to this problem is to analyze the RT data and include them in the analyses or use a drift-diffusion modeling approach to get more precise estimates of memory strength to control for this feature. An alternative is to sub-select participants in each group to have a sample matched on performance (including # of repetitions) and re-run all the analyses in this sub-sample. Without addressing these concerns it is near impossible to interpret the presented data.

Thank you for highlighting this point.

Firstly, we believe that our approach, involving strategic and repeated learning coupled with feedback, enhances the formation of detailed contextual memories. The retrieval procedure also emphasized the need for detailed memory for location. These are critical differences in experimental procedure from previous studies, which enhanced the importance of detailed representations and likely reduced the likelihood of forming gist-like memories.

Indeed, we ceased further learning after the fourth repetition. Extensive piloting, where we initially stopped after the seventh repetition, showed no improvement beyond the fourth repetition. In fact, performance tended to decline due to fatigue. Therefore, we limited the number of repetition cycles to the point where an improvement of performance was still feasible. Even though children exhibited lower final learning performance overall, we believe our procedure facilitated them to reach their maximal performance within the experimental setup.

To address the reviewer’s concern, we included learning data to illustrate the progression of learning (see Fig. 1C, pp. 9-10 in Results).

When interpreting the retention rates, it is important to note that we reported retention rates only for items that were correctly learned (100%) on day 0, day 1, and day 14. This approach meant that different participants had varying numbers of items learned correctly. However, this method enabled us to address our primary question: whether memory consolidation, based on all items initially encoded successfully, is comparably robust between the groups. To simultaneously examine the change in retention rate slopes over time for recent (30 minutes after learning), short delay (one night after) remote, and long delay (two weeks after) remote items, we conducted a separate analysis of retention rates for recent items on days 1 and 14. After observing no differences between sessions in both age groups, we combined the data for recent items. This allowed us to investigate how the slope of memory retention for initially correctly learned items (with a baseline of 100%) changes over time. We observed a significant interaction between item type (recent, short delay remote, long delay remote) and group. Analysis of this interaction revealed significantly less robust memory consolidation across all delay times for children compared to young adults. The figures have been adjusted accordingly to incorporate the baseline of 100% correct performance.

Following your suggestion, we also employed the drift diffusion model approach to characterize memory strength, calculating drift rate, boundary and non-decision time parameters. We added the results to the Supplementary Materials (section S2.1, Figure S1).

Generally, our findings indicate lower overall drift rate in children when considering all items that had to be learned. We also observed that adults show higher slope of decline in drift rate in short and long delay, which, however, are characterized still by higher memory strength compared to children. Both age groups required similar amount of evidence to make decision, which declined with delay. It may indicate an adaptation of weaker memory. Further, we observed lesser non-decision time in children compared to adults, potentially suggesting less error checking or less thorough processing and memory access through strategy in children.

Overall, these results indicate weaker memory strength in children as a quantitative measure. It may nevertheless stem from qualitatively different memory representations that children form, as our RSA findings suggest. We believe that our neural effect reflects the effect of interest (i.e., worse memory due to lower memory strength in children). When controlled for, it will take away variance of interest in the neural data. Therefore, we will refrain from including memory strength into the model. However, we will include mean RT as the indicator of general response tendencies.

Given that the paper is already very complex and long, we opted to add the diffusion model results to the Supplementary Materials (section S2.1, Fig. S1), while discussing the results in the discussion (p. 35).

(2) More discussion of the behavioral task should be included in the results, in particular the nature of the adaptive learning paradigm including the behavioral results as well as the categorical nature of the memoranda. Without this information, it is difficult for the reader to understand what category-level versus item-level reinstatement reflects.

Thank you for this valuable input. We have incorporated this information into the results section. Please refer to pp. 9-10, 12, 14, 21, 25-26 for the added details.

(3) Some of the methods for the reinstatement analysis were unclear to me or warranted further adjustment. I believe the authors compared the scene against all other scenes. I believe it would be more appropriate to only compare this against scenes drawn from the same category as opposed to all scenes. Secondly, from my reading, it seems like the reinstatement was done during the scene presentation, rather than the object presentation in which they would retrieve the scene. I believe the reinstatement results would be much stronger if it was captured during the object presentation rather than the re-presentation of the scene. Or perhaps both sets of analyses should be included.

We apologize for the confusion regarding the analysis method.

During the review process we have improved the description of this analysis and hope it is easier to follow now. In short, we used both approaches (within and between categories) to suit different goals (I.e., measuring scene-reinstatement and gist-like reinstatement).

Both types of reinstatement were assessed during the fixation cross to avoid confounds with the object itself being on the screen. We only used the scene window in one analysis (scene-reinstatement index) as a neural template to track its pre-activation during the fixation. So, as the reviewer suggests, our rationale is that the reinstatement indeed starts taking place at the short object presentation window, but importantly, extends to the fixation window. We added this clarifying information to the results section (see p. 21-27).

(4) For the univariate results, it was unclear to me when reading the results whether they were focusing on the object presentation portion of the trial or the scene presentation portion of the trial. Again, I think the claims of reinstatement related activity would be stronger if they accounted for the object presentation period.

Thank you for pointing this out. Indeed, the univariate results were based on the object presentation time window. We added this information to the results section (Fig. 3, pp. 14, 16).

(5) Further, given the univariate differences shown across age groups, the authors should re-run all analyses for the RSA controlling for mean activation within the ROI.

Thank you for highlighting this. We re-ran all analysis for the RSA controlling for the mean activation within the ROI. The results remained unchanged. We have added this information to the results section as well as in Table S8 and S11 in the Supplementary Materials for further details.

(6) The authors should include explicit tests across groups for their brain-behavior analyses if they want to make any developmentally relevant interpretations of the data. Also, It would be helpful to include similar analyses to those using the univariate signals, and not just the RSA results.

Following reviewer’s suggestion, we included brain-behavior analyses for univariate data as well as RSA data with explicit tests across groups. These can be found in the Results Section pp. 18-20, 28-32. Due to the interdependence of predefined ROIs and to avoid running a high number of correlation tests, we employed the partial least square correlation analysis for this purpose. This approach focuses on multivariate links between specified Regions of Interest (ROIs) and fluctuations in memory performance over short and long delays across different age cohorts. We argue that this multivariate strategy offers a more comprehensive understanding of the relationships between brain metrics across various ROIs and memory performance, given their mutual dependence and connectivity (refer to Genon et al. (2022) for similar discussions).

(7) There could be dramatic differences in memory processing across 5-7 year olds. I know the sample is a little small for this, but I would like to see regressions done within the middle childhood group in addition to the across-group comparisons.

We have included information detailing the relationship between memory retention rate and age within the child group (refer to p. 13). In the child group, both recent and short delay remote memory improved with age. However, the retention rate for long-delayed memory did not show a significant improvement with increasing age in children.

(8) I am concerned that the authors used global-signal as a regressor in their first-level analyses, given that there could be large changes in the amount of univariate activation that occurs across groups. This approach can lead to false positives and negatives that obscure localized differences. The authors should remove this term, and perhaps use the mean sum of the white matter or CSF to achieve the noise regressor they wanted to include.

We understand the reviewers' concerns. However, we believe that our approach is recommended for the pediatric population. Specifically, Graff et al., 2021, found that global signal regression is a highly efficacious denoising technique in their study of 4 to 8-year-old children. This technique was previously suggested for adults by Ciric et al., 2017, and the benefits in terms of motion and physiological noise removal outweigh the potential costs of removing some signal of interest, as indicated by Behzadi et al., 2007. Additionally, we incorporated the six anatomic component-based noise correction (CompCor) to account for WM and CSF signals, as recommended in the pediatric literature.

(9) The authors discuss the relationship between hippocampal reactivation and worse memory through the lens of Schapiro et al., but a new paper by Tanriverdi et al came out in JOCN recently that is more similar to the authors' findings.

Thank you for highlighting the recent paper by Tanriverdi et al. in JOCN, which aligns closely with our findings. We appreciate the suggestion and agree that exploring this alignment could further enrich our discussion on the relationship between hippocampal reactivation and memory retention. We incorporated this work in our revised manuscript .

Minor Comments

- I was surprised that the authors did not see any differences in univariate signals for memory retrieval as a function of development, as much of the prior work has shown differences (for example work by Tracy Riggins). I believe this contrast should be highlighted in the discussion.

- Given the robust differences in sleep patterns across childhood and the role of sleep in systems consolidation framework, I think this feature should be highlighted in either the introduction or discussion.

- Could the authors report on differences (or lack of differences) in head motion across the groups, and if they are different whether they could include them as a confounding variable.

I believe we included six motion parameters and their derivatives into the model

Thank you for your comments.

First, prior works on univariate signals of memory retrieval focused mostly on remembered vs forgotten contrasts, while in our study we focused on remote vs recent in short and long delay only for correctly remembered items. This can partially explain the results. We highlighted this information in the discussion session.

Second, we agree with the reviewer that sleep patterns across childhood should be addressed in the analysis. Therefore, we incorporated them in the discussion section.

Third, indeed head motion were included in the analysis as confounding variables, as adding them is highly recommended for the developmental population (e.g., Graff et al. 2021). As an example, we observed higher framewise displacement in children compared to adults, t = -16(218), p <. 001, as well as in translational y, t = -2.33(288), p = .02.

Reviewer #3 (Public Reviews):

Summary:

This study aimed to understand the neural correlates of memory recall over short (1-day) and long (14-days) intervals in children (5-7 years old) relative to young adults. The results show that children recall less than young adults and that this is accompanied by less activation (relative to young adults) in brain networks associated with memory retrieval.

Strengths:

This paper is one of few investigating long-term memory (multiple days) in a developmental population, an important gap in the field. Also, the authors apply a representational similarity analysis to understand how specific memories evolve over time. This analysis shows how the specificity of memories decreases over time in children relative to adults. This is an interesting finding.

We thank the reviewer for the appraisal of our manuscript.

Weaknesses:

Overall, these results are consistent with what we already know: recall is worse in children relative to adults (e.g., Cycowicz et al., 2001) and children activate memory retrieval networks to a lesser extent than adults (Bauer et al, 2017).

It seems that the reduced activation in memory recall networks is likely associated with less depth of memory encoding in children due to inattentiveness, reduced motivation, and documented differences in memory strategies. In regard to this, there was consideration of IQ, sex, and handedness but these were not included as covariates as they were not significant although I note p<.16 suggests there was some level of association nonetheless. Also, IQ is measured differently for the children and adults so it's not clear these can be directly contrasted. The authors suggest the instructed elaborative encoding strategy is effective for children and adults but the reference in support of this (Craik & Tulving, 1975) does not seem to support this point.

Thank you for your review, and we appreciate your valuable feedback. Here are our responses and clarifications:

Regarding the novelty of the results in terms of mentioned existent literature, we believe that in contrast to Cycowicz et al. (2001) and Bauer et al (2017), etc, we assess not only immediate memory after encoding with semantic judgement of abstract associations, but add to these findings investigating consolidation-related changes in complex associative and contextual information in much under investigated sample of 5-to-7-year-old preschoolers. With this we are able to infer also how neural representations of children change over time, providing invaluable insights into knowledge formation in this developmental cohort.

With this, the observed age differences are not so of primary importance, as time-related changes in mnemonic representations observed in children.

Regarding the assumption of inattentiveness in children, we want to emphasize that the experimenter was present throughout the learning process, closely supervising the children. We observed prompt responses to every trial in children and noted an increase in accuracy over the encoding-learning cycles, leading us to conclude that the children were indeed attentive to the task. The observed accuracy improvement across learning cycles indicates increase in remembered information. Furthermore, we took measures to ensure their engagement, including extensive training in both verbal and computerized versions to ensure that they understood and actively created stories to support their learning.

We collected motivation data after each task execution in children, and the results indicated that they scored high in motivation. Children not only completed the tasks but also expressed their willingness to participate in subsequent appointments, highlighting their active involvement in the study.

The observed differences in the efficiency of strategy utilization were expected, given developmental differences in the associative and strategic components of memory in children, as noted in prior research (Shing, 2008, 2010).

We appreciate your point about IQ, sex, and handedness. These variables were indeed included in the behavioral models, and mean brain activation was also included in the brain data models, addressing the potential influence of these factors on our results.

While it's true that we applied different tests to measure IQ in children and adults, these tests targeted comparable subtests that addressed similar cognitive constructs. As the final IQ values are standardized, we believe it is appropriate to compare them between the two groups.

Lastly, we agree that the citation Craik & Tulving, 1975 supports the notion of effectiveness of instructed elaborative learning only in adults, but not in children. For this purpose, we added relevant literature for the child cohort (i.e., Pressley, 1982; Pressley et al., 1981; Shing et al., 2008).

Reviewer #3 (Recommendations For The Authors):

An additional point for the authors to consider is that the hypotheses were uncertain. The first is that prefrontal, parietal, cerebellar, occipital, and PHG brain regions would have greater activation over time in adults and not children - which is very imprecise as this is basically the whole brain. Moreover, brain imaging data may be in opposition to this prediction: e.g., the hippocampus has a delayed maturational pattern beyond 5-yrs (e.ge., Canada 2019; Uematsu 2012) and some cortical data predicts earlier development in these regions.

Thank you for your feedback, and we appreciate your insights regarding our hypotheses.

The selection of our regions of interest (ROIs) was guided by prior literature that has demonstrated the interactive involvement of multiple brain areas in memory retrieval and consolidation processes. Additionally, our recent work utilizing multivariate partial least square correlation analysis (Schommartz, 2022, Developmental Cognitive Neuroscience) has indicated that unique profiles derived from the structural integrity of multiple brain regions are differentially related to short and long-delay memory consolidation.

Indeed, the literature suggests that the hippocampus may exhibit a more delayed maturational pattern extending into adolescence, as supported by studies such as Canada (2019) and Uematsu (2012), etc. We added this information as well as findings from the literature on cortical development to be more balanced in our review of the literature.

Given this complexity, we believe it is important to emphasize in our discussion that both the medial temporal lobe, including the hippocampus, and cortical structures, as well as the cerebellum, undergo profound neural maturation. We highlight these nuances in our revised manuscript to provide a more comprehensive perspective on the developmental differences in memory retention over time.

The writing was challenging to follow - consider as an example on page 9 the sentence that spans 10 lines of text.

Thank you for bringing this to our attention. We have carefully reviewed the manuscript and have made efforts to streamline the text, ensuring that sentences are not overly long or complex to improve readability and comprehension.

I found the analysis (and accompanying figures) a bit of a data mine - there are so many results that are hard to digest and in other cases highly redundant one from the other. This may be resolved in part by moving redundant findings to the supplemental. Some were hard to follow - so when there is a line between recent and recent data, that seems confusing to connect data that, I believe, are different sets of items. Later scatterplots (Fig 7) have pale yellow dots that I had a hard time seeing.

Thank you for bringing up your concerns regarding the analysis and figures in our manuscript. We have carefully considered your feedback and made several improvements to address these issues.

To alleviate the challenge of digesting numerous results, we have taken steps to enhance clarity and reduce redundancy. Specifically, we have moved some of the redundant findings to the supplementary sections, which should help streamline the main manuscript and make it more reader friendly.

Regarding the line between 'recent' and 'recent data,' figure were transformed to a clearer version. Furthermore, we have improved the visibility of certain elements, such as the pale-yellow dots in the scatterplots (Fig 1, 2, 4, etc. ), to ensure that readers can better discern the data points.

Changes from the earlier version to the revised version

Abstract

Abstract

Added in the revised version: Research highlights

Added in the revised version: Children showed less robust memory consolidation across short and long delay compared to young adults.

Added in the revised version: From short to long delay, children show differential neural upregulation for remote versus recent memory compared to young adults.

Added in the revised version: Over time, both children and young adults showed reduced scene-specific reinstatement of neural patterns.

Added in the revised version: Children relied more on gist-like neural reinstatement in anterior hippocampal and medial prefrontal brain regions.

Earlier version: Memory consolidation tends to be less robust in childhood than adulthood. However, little is known about the corresponding functional differences in the developing brain that may underlie age-related differences in retention of memories over time. This study examined system-level memory consolidation of object-scene associations after learning (immediate delay), one night of sleep (short delay), as well as two weeks (long delay) in 5-to-7-year-old children (n = 49) and in young adults (n = 39), as a reference group with mature consolidation systems. Particularly, we characterized how functional neural activation and reinstatement of neural patterns change over time, assessed by functional magnetic resonance imaging combined with representational (dis)similarity analysis (RSA). Our results showed that memory consolidation in children was less robust (i.e., more forgetting) compared to young adults. For correctly retained remote memories, young adults showed increased neural activation from short to long delay in neocortical (parietal, prefrontal and occipital) and cerebellar brain regions, while children showed increased neural activation in prefrontal and decrease in neural activity in parietal brain regions over time. In addition, there was an overall attenuated scene-specific memory reinstatement of neural patterns in children compared to young adults. At the same time, we observed category-based reinstatement in medial-temporal, neocortical (prefrontal and parietal), and cerebellar brain regions only in children. Taken together, 5-to-7-year-old children, compared to young adults, show less robust memory consolidation, possibly due to difficulties in engaging in differentiated neural reinstatement in neocortical mnemonic regions during retrieval of remote memories, coupled with relying more on gist-like, category-based neural reinstatement.

Revised version: Memory consolidation tends to be less robust in childhood than adulthood. However, little is known about the corresponding functional differences in the developing brain that may underlie age-related differences in retention of memories over time. This study examined system-level memory consolidation of object-scene associations after learning (immediate delay), one night of sleep (short delay), as well as two weeks (long delay) in 5-to-7-year-old children (n = 49) and in young adults (n = 39), as a reference group with mature consolidation systems. Particularly, we characterized how functional neural activation and reinstatement of neural patterns change over time, assessed by functional magnetic resonance imaging combined with representational similarity analysis (RSA). Our results showed that memory consolidation in children was less robust and strong (i.e., more forgetting) compared to young adults. Contrasting correctly retained remote versus recent memories across time delay, children showed less upregulation in posterior parahippocampal gyrus, lateral occipital cortex, and cerebellum than adults. In addition, both children and adults showed decrease in scene-specific neural reinstatement over time, indicating time-related decay of detailed differentiated memories. At the same time, we observed more generic gist-like neural reinstatement in medial-temporal and prefrontal brain regions uniquely in children, indicating qualitative difference in memory trace in children. Taken together, 5-to-7-year-old children, compared to young adults, show less robust memory consolidation, possibly due to difficulties in engaging in differentiated neural reinstatement in neocortical mnemonic regions during retrieval of remote memories, coupled with relying more on gist-like generic neural reinstatement.

Introduction

Every day we form new memories that may become long-lasting through memory consolidation, a complex process in flux between encoding and retrieval (Dudai, 2012; Josselyn et al., 2015; Moscovitch & Gilboa, 2022; Semon, 1921).During systems-level consolidation, memory representations and traces are reorganized across medial temporal lobe and neocortical brain networks (Ranganath & Ritchey, 2012; Ritchey & Cooper, 2020). These networks include brain regions that are involved both in initial encoding and in integration of new memories as time passes (Axmacher & Rasch, 2017; Dudai, 2012; Moscovitch & Gilboa, 2022; Squire et al., 2015). While decades of work have shed light on general neural mechanisms of memory consolidation in adults (Moscovitch & Gilboa, 2022; Sekeres et al., 2017a; Winocur & Moscovitch, 2011), much less is known about neural mechanisms that support memory consolidation in children – a knowledge gap that we aimed to address with the current study.

Neural correlates of memory consolidation

Learning through repeated activation and reinstatement is one way to rapidly stabilize memory traces and make them accessible upon retrieval (Dudai, 2004; Nader & Hardt, 2009; Teyler & Rudy, 2007). For instance, in young adults, repeated exposure to word-image pairs during encoding, compared to single exposure, was shown to accelerate memory consolidation. This is achieved through enhanced replay of repeated events in the retrosplenial cortex (RSC) and the medial prefrontal cortex (PFC), as well as via increased hippocampal (HC)-cortical replay that promotes the associative word-object memories (Yu et al., 2022). In another study by Brodt et al. (2016), it was found that during repeated spatial navigation in a virtual environment, activation in the posterior parietal cortex (PPC), especially the precuneus, increased and remained elevated after 24 hours, while HC activity and HC-PPC connectivity declined with repeated encoding rounds (Brodt et al., 2016). In addition, neocortical plasticity measured by diffusion-weighted magnetic resonance imaging in the PPC (Brodt et al., 2018) and the cerebellum (Stroukov et al., 2022) supported rapid cortical storage of memory traces for object-location associations after repeated exposure in young adults 1 hour and 12 hours post-learning. Taken together, these findings indicate that repeated learning in young adults promotes fast creation of neural memory representations, which can remain stable for at least 24 hours and predict behavioural mnemonic performance.

Earlier version: Memory consolidation of well-learnt information does not end with the last learning cycle, but undergoes further neural reorganizing and modification over time (Roüast & Schönauer, 2023; Sekeres et al., 2017). For example, during cued recall of face-location associations, young adults who were tested 24 hours after learning, compared to 15 minutes, showed increased activation in the precuneus, inferior frontal gyrus (IFG), and fusiform gyrus, whereas the hippocampus showed a decrease in activation (Takashima et al., 2009). Similarly, increased activation in the anterior temporal cortex during the retrieval of studied figure pairs eight weeks prior was observed, while increased activation in the HC was shown for pairs learned immediately before retrieval (Yamashita et al., 2009). Furthermore, delayed retrieval of naturalistic video clips after the delay of seven days in young adults was associated with the increased activation in the lateral and medial PFC and decrease in HC and parahippocampal (PHG) activation over time (Sekeres et al., 2021). This is convergent with the notion that the role of the prefrontal cortex increases during recollection as consolidation progresses over time (Milton et al., 2011). Moreover, subsequently recollected memories showed higher post-rest HC-lateral occipital cortex (LOC) connectivity specifically related to scene-related mnemonic content, indicating the role of LOC in associative memory consolidation (Tambini et al., 2010).

Revised version: Memory consolidation of well-learnt information does not end with the last learning cycle, but undergoes further neural reorganizing and modification over time (Roüast & Schönauer, 2023; Sekeres et al., 2017). For example, during cued recall of face-location associations, young adults who were tested 24 hours after learning, compared to 15 minutes, showed increased activation in the precuneus, inferior frontal gyrus (IFG), and fusiform gyrus, whereas the hippocampus showed a decrease in activation (Takashima et al., 2009). Similarly, increased activation in the anterior temporal cortex during the retrieval of studied figure pairs eight weeks prior was observed, while increased activation in the HC was shown for pairs learned immediately before retrieval (Yamashita et al., 2009). Furthermore, delayed retrieval of naturalistic video clips after the delay of seven days in young adults was associated with the increased activation in the lateral and medial PFC and decrease in HC and parahippocampal (PHG) activation over time (Sekeres et al., 2021). This is convergent with the notion that the role of the prefrontal cortex increases during recollection as consolidation progresses over time (Milton et al., 2011). Moreover, subsequently recollected memories showed higher post-rest HC-lateral occipital cortex (LOC) connectivity specifically related to scene-related mnemonic content, indicating the role of LOC in associative memory consolidation (Tambini et al., 2010). On the other hand, HC activation has been reported to remain stable after seven days (Sekeres, Winocur, Moscovitch, et al., 2018) three months (Harand et al., 2012) or even years (Söderlund et al., 2012) for consistent episodic memories that retained contextual details.

Earlier version: To summarize, studies have shown that with passing time, memories of well-learned information increasingly engage cortical regions including the prefrontal, parietal, occipital, and anterior temporal brain areas that support retrieval of general and schematic memories, as well as complex associative information, while the recruitment of HC tends to decrease. However, most research focused on only a selected delay window and only on young adults.

Revised version: To summarize, in alignment with the Multiple Trace Theory (Nadel et al., 2000; Nadel & Moscovitch, 1997), studies have shown that memories of well-learned information increasingly engage cortical regions over time. There regions include the prefrontal, parietal, occipital, and anterior temporal brain areas, supporting the retrieval of general and schematic memories, as well as complex associative information. In line with the Standard Consolidation Theory, some studies have demonstrated a decrease in the recruitment of the HC over time (Squire & Alvarez, 1995). Conversely, and converging with the Contextual Binding Theory (Yonelinas et al., 2019) and the Multiple Trace Theory, some studies have shown that hippocampal involvement lingers over time, particularly for detailed and contextual memories. However, most research has focused on only a selected delay window and solely on young adults.

Earlier version: Mnemonic reinstatement across consolidation

Revised version: Mnemonic transformation and reinstatement across consolidation

Earlier version: In addition to changes in neural activation during mnemonic retrieval over time, it is important to characterize the transformations of neural representations (i.e., distinctive pattern of neural activity generated by a specific memory; Averbeck et al., 2006; Kriegeskorte, 2008; Kriegeskorte & Kievit, 2013) because the multivariate activity pattern of memory may change over time. For example, memory for perceptual details may become worse over time, while memory for gist may be more likely to stay stable (Reyna & Brainerd, 1995; Sekeres et al., 2016). Little is known about how the neural representation of well-learned memories at retrieval is transformed across the consolidation period (i.e., phenomenon, when similar patterns of neural activity may be reactivated when memory is retrieved again; Clarke et al., 2022; Deng et al., 2021). Using representational similarity analysis (RSA; Kriegeskorte, 2008), Tompary & Davachi (2017) showed that a one-week delay led to memory reorganisation in HC and mPFC. Specifically, during a one-week delayed retrieval compared to immediate retrieval the overlap in neural patterns reflected the overlap in stimulus features. Comparing neural reinstatement of visual clips during encoding, immediate, and delayed recall (after 1-week-period), Oedekoven et al. (2017) showed reliable reinstatement in core retrieval networks, including the precuneus, medial temporal gyrus, occipital gyrus, HC, and PHG among others. In contrast to Tompary and Davachi (2018), this study found no time-related differences in reinstatement effects. Therefore, the findings on memory reinstatement are mixed, and, to date, no study have directly tracked the neural representations of memory traces for perceptual together with more abstract, gist-like features (e.g., semantic categories).

Revised version: In addition to changes in neural activation during mnemonic retrieval over time, it is important to characterize the transformations and reinstatement of neural representations (i.e., distinctive pattern of neural activity generated by a specific memory; Averbeck et al., 2006; Kriegeskorte, 2008; Kriegeskorte & Kievit, 2013) because the multivariate activity pattern of memory may change over time. For example, memory for perceptual details may become worse over time, while memory for gist may be more likely to stay stable, indicating differential time-related transformational trajectories (Sekeres et al., 2016). According to the Fuzzy Trace Theory (Reyna & Brainerd, 1995, 1998), detailed and gist-like memories may be uniquely present or coexist, depending on the strength of formed memories. For instance, detailed memories may generally fade away over time, preserving however its specific accurate nature for correctly recalled memories (Diamond et al., 2020). In other instances, weaker detailed memories may be reorganized over time, with lingering specific memories and parallel creation of gist-like generic memories. Little is known about how the neural representation of well-learned memories at retrieval is transformed across the consolidation period (i.e., phenomenon, when similar patterns of neural activity may be reactivated when memory is retrieved again; Clarke et al., 2022; Deng et al., 2021).

Added in the revised version: Using representational similarity analysis (RSA; Kriegeskorte, 2008), Tompary & Davachi (2017) showed that a one-week delay led to differential memory reorganisation in HC and mPFC for memories with and without overlapping features. Specifically, after a one-week mnemonic representations became more similar for memories with overlapping features, indicating consolidation-related gist-like neural reorganization. Moreover, the authors showed memory-specific reinstatement of neural patterns for specific memories in the right HC, indicated by significant encoding-retrieval similarity for remote but not recent memories. Comparing neural reinstatement of visual clips during encoding, immediate, and delayed recall (after 1-week-period), Oedekoven et al. (2017) showed reliable reinstatement in core retrieval networks, including the precuneus, medial temporal gyrus, occipital gyrus, HC, and PHG among others. In contrast to Tompary and Davachi (2018), this study found no time-related differences in reinstatement effects. Therefore, the findings on memory reinstatement are mixed, and, to date, no study have directly tracked the neural representations of memory traces for perceptual together with more abstract, gist-like features (e.g., semantic categories).

Earlier version: Neural correlates of memory consolidation and mnemonic reinstatement in middle childhood

Revised version: Neural correlates of memory consolidation and mnemonic transformation and reinstatement in middle childhood

Earlier version: Brain regions involved in memory consolidation show protracted developmental trajectories from early to late childhood (Badre & Wagner, 2007b; Ghetti & Bunge, 2012c; Gogtay et al., 2004; Keresztes et al., 2022; Lenroot & Giedd, 2006; Mills et al., 2016; Ofen et al., 2007; Shing et al., 2008), which could lead to differences in neural activity and/or patterns of mnemonic reinstatement between children and adults. For instance, univariate selectivity was reduced in children, while fine-grained representational similarity along the ventral visual stream for mnemonic reinstatement was similar in 5-11 years old children and adults (Cohen et al., 2019; Golarai et al., 2015). These findings indicate that mnemonic reinstatement may develop prior to univariate selectivity. Fandakova et al. (2019) also showed that neural specificity during encoding was similar in 8-to-15-year-old children and adults in the RSC, LOC and PHG. Neural specificity was also associated with subsequent memory in a similar way between children and adults. However, it is unclear whether the age-related differences in neural activation and reinstatement mentioned above are similar for memory consolidation. Specifically, to what extent does consolidation-related transformation of neural representations occur, and how does it impact neural reinstatement of mnemonic content in the developing brain?

Revised version: Brain regions involved in memory consolidation show protracted developmental trajectories from early to late childhood (Badre & Wagner, 2007b; Ghetti & Bunge, 2012c; Gogtay et al., 2004; Keresztes et al., 2022; Lenroot & Giedd, 2006; Mills et al., 2016; Ofen et al., 2007; Shing et al., 2008), which could lead to differences in neural activity and/or patterns and subsequently mnemonic reinstatement between children and adults. For instance, univariate selectivity was reduced in children, while fine-grained neural representational similarity along the ventral visual stream was similar in 5-11 years old children and adults (Cohen et al., 2019; Golarai et al., 2015). Fandakova et al. (2019) also showed that the neural representational distinctiveness of information during encoding was similar in 8-to-15-year-old children and adults in the RSC, LOC and PHG. The fidelity of neural representations was also associated with subsequent memory in a similar way between children and adults. Overall, although these findings did not address the question of neural reinstatement directly in children, they suggest that mnemonic reinstatement may develop prior to univariate selectivity. However, it is yet to be investigated. Moreover, it is unclear whether the age-related differences in neural activation and reinstatement mentioned above are similar for memory consolidation. Specifically, to what extent does consolidation-related transformation of neural representations occur, and how does it impact neural reinstatement of mnemonic content in the developing brain?.

In middle childhood, the trade-off between retaining vivid, detail-rich memories and their transformation into vague, gist-like memories due to delay may be more pronounced. Brainerd et al., (2002) demonstrated that, during development, specific memory and gist-memory for events emerge together. However, as children mature, they exhibit more false memories based on gist in the absence of exact memories for the events. On the other hand, Keresztes et al. (2018) postulated that younger children tend to rely more on generalization when forming new memories, while older children and adults use more specific detail-rich information, suggesting a shift from generalization to specificity as children mature. Hence, there are some inconsistencies in the theoretical postulations and findings regarding item-specific and gist-based memories that may impact memory consolidation in middle childhood. Investigation on the neural reinstatement patterns of item-specific and gist-like memories across time may add to the understanding of these inconsistencies in children.

Aim of the current study

Earlier version: In this study, we examined the univariate neural activation and multivariate neural reinstatement patterns of memories for object-location associations across a short delay (after one night of sleep) and a long delay (after a 2-week period), relative to recently consolidated memories (after 30 minutes). Children (5-to-7-year-old) were compared to young adults serving as a reference group with a mature memory consolidation system. We selected 6 to 7 years as the age range of interest because previous studies showed a large improvement in associative memory around this age (Sluzenski, Newcombe, & Kovacs, 2006). Practically, this is also the youngest age range in which MRI scanning coupled with active task execution could be applied relatively successfully. We hypothesized (i) an increasing involvement of prefrontal, parietal, cerebellar, occipital and PHG brain regions over time in adults in comparison to children, as these regions are still maturing in preschool and early school-aged children (Ghetti & Bunge, 2012a; Keresztes et al., 2022; Lebel et al., 2012; Shing et al., 2008, 2010a); (ii) a stable involvement of HC over time in adults and children due to relative maturity of the HC in middle childhood (Keresztes et al., 2017; Sekeres et al., 2018; Shing et al., 2008; Sluzenski et al., 2006); (iii) a decreasing neural reinstatement in all ROIs over time, with this decrease being more pronounced in children compared to young adults (Cohen et al., 2019; Golarai et al., 2015); (iv) different contributions of category- and item-specific memories to neural reinstatement across age groups. Specifically, we expected more gist category-based memory pattern reinstatement in children in comparison to more detailed item-specific neural pattern reinstatement in young adults over time (Reyna & Brainerd, 1995)).

Revised version: In this study, we examined the univariate neural activation and multivariate neural reinstatement patterns of memories for object-location associations across a short delay (after one night of sleep) and a long delay (after a 2-week period), relative to recently consolidated memories (after 30 minutes). Children (5-to-7-year-old) were compared to young adults serving as a reference group with a mature memory consolidation system. We selected 5 to 7 years as the age range of interest because previous studies showed a large improvement in associative memory around this age (Sluzenski, Newcombe, & Kovacs, 2006). Practically, this is also the youngest age range in which MRI scanning coupled with active task execution could be applied relatively successfully. We hypothesized (i) according to the Multiple Trace Theory, an increasing involvement of prefrontal, parietal, cerebellar, occipital and PHG brain regions over time in adults in comparison to children, as these regions are still maturing in preschool and early school-aged children (Ghetti & Bunge, 2012b; Keresztes et al., 2022; Lebel et al., 2012; Shing et al., 2008, 2010a); (ii) according to the Contextual Binding Theory, the Multipe Trace Theory, and supported by the evidence from Sekeres, Winocur, & Moscovitch (2018), a stable involvement of HC over time in adults and children due to relative maturity of the HC in middle childhood and detailed contextual nature of the repeatedly learned information (Keresztes et al., 2017; Nadel et al., 2000; Sekeres, Winocur, Moscovitch, et al., 2018; Shing et al., 2008; Sluzenski et al., 2006; Yonelinas et al., 2019); (iii) a decreasing neural reinstatement in all ROIs over time, with this decrease being more pronounced in children compared to young adults (Cohen et al., 2019; Golarai et al., 2015); (iv) different contributions of category- and item-specific memories to neural reinstatement across age groups. Specifically, we expected more gist-like memory pattern reinstatement in children in comparison to more detailed item-specific neural pattern reinstatement in young adults over time due to differences in the strength of formed memories and differences in underlying associative and strategic components of memories (Reyna & Brainerd, 1995; Shing et al., 2008, 2010).This assumption aligns with the Fuzzy Trace Theory (Brainerd & Reyna, 2002) that verbatim memories can be created without the extraction of gist. Due to ongoing maturation of associative and strategic memory components and their underlying neural substrates, children may be more inclined to extract gist information at the expense of detailed or gist-like information.

Results

Behavioural results

Final Learning Performance

Added in the revised version: Unique sets of object-location association pairs were learned on Day 0, Day 1, and Day 14. During each initial encoding trial, participants were presented with an object within a congruent scene (e.g., a fox in a spring pine tree forest), and were asked to memorize the exact location on the object within the scene by creating a story and making “mental” pictures of the scene. The choices for locations varied across scenes while they remained constant across time within individuals. There were 18 unique key locations among which object could be distributed, resulting in a heterogenous set of locations for objects. We employed an adaptive, repetitive learning-to-criteria procedure to ensure initially strong memories (see Fig. 1A for the task overview and Fig. 1B for experimental procedure overview). Before learning began, participants were instructed to create stories to help them memorize the locations of the objects within the scenes. They practiced this strategy on two unique sets of five object-location associations. Subsequently, the learning began with the first encoding block. Following each encoding block, the repetitive learning-to-criteria started. During learning, participants were presented with the scenes again, but with three rectangles indicating possible locations for the previously learned objects. The choice options for the three alternative forced choice task (3AFC) were distributed variably: for some instances, the “correct” answer was the left option, for some instance – it was the middle option, and for other instances it was the right option. Therefore, the correct performance in the task required detailed and precise memory for locations. Moreover, the choice options were presented rather close together in the scene, requiring higher level of differentiation among the options. Participants were asked to choose one rectangle that corresponded to the correct location of the object within the scene (Fig 1A “Learning Cycles”). Regardless of accuracy, the object was shown again in the correct location. The learning procedure was repeated at least two times and maximally four times or until the response accuracy of 83% was reached within one cycle.

Added in the revised version: (A) Trial Structures in the Experimental Task.

Removed in the revised version: For final learning accuracy, operationalized as percentage of correctly identified locations relative to the total number of locations (see Fig. 1A for the task overview and Fig. 1B for experimental procedure overview), the LME model revealed a significant effect of Group, F(1,79) = 94.31, p < .001, w2 = .53, showing higher overall final accuracy in young adults in comparison to children t(185) = 7.55, p < .001 (Fig. 2A). No Session effect (p = .79) or Session x Group interaction was significant (p = .96), indicating a stable level of final learning accuracy in each age group across sessions with different stimuli sets. Although the learning procedure was adaptive, the memory performance of children was inferior to that of young adults at the end of learning.

Removed in the revised version: Memory Task and Experimental Procedure

Earlier version: (A) Trial Structures in the Experimental Task. (i) In the Initial Encoding phase, participants were instructed to remember object-location pairs by creating a story or making a “mental photo” of the scene, memorizing the exact location of each object within the scene. (ii) In the Learning Phase, participants chose one location out of three choices and received feedback for their response. After receiving feedback, the correct object-location association was shown again. (iii) In the Retrieval Phase participants chose the location of the object in the scene out of three options without feedback. The retrieval phase took place in the MR scanner. (B) Experimental Procedure. The testing took place across three days. On Day 0, participants learned 60 object-location associations (remote items). On Day 1, participants learned 30 new object-location associations (recent items). For retrieval (short delay), 30 remote pairs learned on Day 0 and 30 recent pairs learned on Day 1 were retrieved. A similar procedure was followed on Day 14 (long delay), with another 30 new object-location associations. Across all testing days, participants also completed socio-demographic questionnaires and other psychometric tests, which were distributed across sessions. Note: RT – reaction time; s – second, fMRI – functional magnetic resonance imaging.

Revised version: (i) In the Initial Encoding phase, participants were instructed to remember object-location pairs by creating a story or making a “mental photo” of the scene, memorizin the exact location of each object within the scene. (ii) In the Learning Phase, participants chose on location out of three choices and received feedback for their response. The feedback was given in the form of a smiley face, with a happy face denoting a correct answer, a sad face denoting an incorrect answer, an a sleeping face denoting a missed response. After receiving feedback, the correct object-locatio association was shown again. (iii) In the Retrieval Phase participants chose the location of the object in the scene out of three options without feedback. The retrieval phase took place in the MR scanner. (B) Experimental Procedure. The testing took place across three days. On Day 0, participants learned 60 object-location associations (remote items). On Day 1, participants learned 30 new object-location associations (recent items). For retrieval (short delay), 30 remote pairs learned on Day 0 and 30 recent pairs learned on Day 1 were retrieved. A similar procedure was followed on Day 14 (long delay), with another 30 new object-location associations. Across all testing days, participants also completed socio-demographic questionnaires and other psychometric tests, which were distributed across sessions. Note: RT – reaction time; s – second, fMRI – functional magnetic resonance imaging. (C) Overview of Learning Performance. Children needed on average between two to four learning-retrieval cycles to reach the criterion of 83% correct responses, while young adults needed on average two cycles. (D) Final Learning Performance. Final learning accuracy is calculated as the percentage of correct responses during the last learning cycle for both children and young adults. Final learning accuracy was significantly higher in young adults compared to children across all sessions. Grey dashed line indicates the criteria of 83% correctly learned items.

Added in the revised version: Concerning number of learning cycles, the linear mixed effects (LME) model revealed a significant Group effect, F(1,563) = 7.09, p. 008, w2 = .01, with children needing more learning cycles to reach the learning criteria in comparison to adults, b = -.43, t(563) = −2.66, p = .008. On average, children needed between two to four learning-retrieval cycles to reach the criterion of 83% correct responses, while young adults required on average two cycles (Fig. 1C). The number of learning cycles did not differ between sessions as revealed by non-significant Session effect and Group x Session interaction (all p > .40).

Added in the revised version: For final learning accuracy, operationalized as percentage of correctly identified locations relative to the total number of locations, the LME model revealed a significant effect of Group, F(1,79) = 94.31, p < .001, w2 = .53, showing higher overall final accuracy in young adults in comparison to children t(185) = 7.55, p < .001 (Fig. 1D). No Session effect (p = .79) or Session x Group interaction was significant (p = .96), indicating a stable level of final learning accuracy in each age group across sessions with different stimuli sets. Although the learning procedure wa adaptive, the memory performance of children was inferior to that of young adults at the end of learning.

Removed in the revised version: Overview if Mnemonic Performance

Removed in the revised version: (A) Learning Performance. Final learning accuracy is calculated as the percentage of correct responses during the last learning cycle for both children and young adults. Final learning accuracy was significantly higher in young adults compared to children across all sessions. Children needed between two to four learning-retrieval cycles to reach the criterion of 83% correct responses, while young adults required on average two cycles. Grey dashed line indicates the criteria of 83% correctly learned items. (B) Memory retrieval. Memory accuracy is operationalized as the percentage of correct responses in the retrieval task conducted during the MRI scanning sessions. Memory accuracy for recently consolidated items did not differ between sessions in young adults, while for children, recent memory accuracy on Day 1 was higher than on Day 14. Memory accuracy for remotely consolidated items differed between sessions in both young adults and children, showing higher remote memory accuracy on Day 1 than on Day 14. All tests used Sidak correction for multiple comparisons. Red dashed line indicates the threshold for random performance. *p < .05; **p < .01; ***p < .001(significant difference); non-significant differences were not specifically highlighted. Error bars indicate standard error based on the underlying LME-model.

Memory Retention Across Time

Added in the revised version: Change in memory retention was investigated during the retrieval part of the memory task (Fig. 1A “Retrieval (fMRI)”). Participants were cued with the object and were instructed to recall as vividly as possible the associated scene and location of the object within the scene during the fixation time window, where no visual input was presented on the screen. After that the associated scene with three choice options was presented and participants had to choose one rectangle denoting the correct location of the object in the scene (see Methods for more details).

Added in the revised version: First, we investigated whether retention rates for recent items initially correctly learned on Day 1 and Day 14 differed between sessions in children and adults. We observed no significant Session x Group interaction, F(1,75) = 1.77, p =.187, w2 = .001, indicating that the difference between retention rates for recent items on Day 1 and Day 14 for initially correctly learned items did not significantly differ between children and young adults. Based on that, we collapsed recent retention rates across sessions in each group for the further analysis.

Added in the revised version: Second, we examined change in memory retention rates for items that were initially correctly learned (i.e., initially strong memories) particularly testing for group differences in recent and remote (short- and long-delay) memory retention in relation to baseline of 100% (see Fig. 2, Supplementary Table S1 for a full overview). The linear mixed-effects model for retrieval accuracies of learned object-location pairs explained a significant amount of variance, R2 = .77, 95% CI [.73 – .81]. We observed a significant main effects of Item Type, F(3,250) = 229.18, p <.001, w2 = .73, indicating overall no difference between recent memory retention compared to short delay remote memory retention, b = 1.49, t(259) = 1.26, p = .754, but higher recent memory retention compared to long delay remote memory retention, b = 21.36, t(259) = 17.59, p < .001, and higher short delay remote memory retention compared to long delay remote memory retention, b = 19.88, t(260) = 16.16, p < .00. Further, we observed a significant main effect of Group, F(1,85) = 55.00, p <.001, w2 = .38, indicating overall lower memory retention in children compared to young adults, b = −11.1, t(91) = −7.20, p < .001. Additionally, we observed a significant Item Type x Group interaction, F(3,250) = 17.35, p < .001, w2 = .16. Model-based Sidak post hoc comparisons revealed that in children group there was a significant decline in memory retention rates for correctly learned recent items, b = 17.18, t(254) = 11.09, p < .001, short delay remote items, b = 16.74, t(255) = 10.60, p < .001, and long delay remote items, b = 37.45, t(260) = 22.87, p < .001. In young adults’ group, there was no significant decline in memory retention rates for correctly learned recent items, b = 1.91, t(254) = 1.10, p = .983, but for short delay remote items, b = 5.32, t(254) = 3.05, p = .033, and long delay remote items, b = 24.37, t(258) = 13.58, p < .001. Additionally, the slope of memory retention decline was significantly steeper in children compared to adults for recent items, b = 15.26, t(254) = 6.56, p < .001, for short delay remote items, b = 11.41, t(255) = 4.84, p < .001, and for long delay remote items, b = 13.08, t(258) = 5.38, p < .001. Furthermore, we observed that memory retention rates significantly increased with age in the child group for recent items, b = .89, t = 2.62, p = .016(FDR-corrected), for short delay remote items, b = .91, t = 2.67, p = .016(FDR-corrected), but not for long delay remote items, b = .15, t = .326, p = .747(FDR-corrected).

Added in the revised version: Retention rates for initially correctly learned items.

Added in the revised version: Memory accuracy is operationalized as the percentage of correct responses in the retrieval task conducted during the MRI scanning sessions for items that were initially correctly learned, indicating initially strong memories. Memory accuracy for recently consolidated items did not differ between sessions in young adults and children and was collapsed across recent memory accuracy on Day was higher than on Day 14. Memory accuracy for remotely consolidated items differed between sessions in both young adults and children, showing higher remote memory accuracy on Day 1 than on Day 14. All tests used Sidak correction for multiple comparisons. Red dashed line indicates the threshold for random performance. *p < .05; **p < .01; ***p < .001(significant difference); non-significant differences were not specifically highlighted. Error bars indicate standard error based on the underlying LME-model.

Added in the revised version: Taken together, both age groups showed a decline in memory performance over time. However, compared to young adults, children showed a steeper slope of memory decline for both immediate recent and remote short and long delay memories. In sum, the results showed that children had overall worse memory retention rates compared young adults, indicating less robust memory consolidation in children.

Removed in the revised version: In the following, we examined the change in memory performance for correctly learned items across time (Fig. 2B, Table S1-S2 for full overview). We observed significant main effects of Session, F(1,241) = 150.31, p <.001, w2 = .38, indicating higher overall memory retention on Day 1 compared to Day 14, b = 10.9, t(250) = 12.09, p < .001; of Item Type, F(1,229) = 203.25, p <.001, w2 = .47, indicating higher overall memory retention for recent compared to remote items, b = 12.3, t(238) = 14.08, p < .001; of Group, F(1,84) = 70.13, p <.001, w2 = .45, indicating lower overall memory retention in children compared to young adults, b = -18, t(93) = -8.08, p < .001; and Session x Item type interaction, F(1,229) = 80.96, p < .001, w2 = .26. Model-based Sidak post hoc comparisons revealed that memory retention across both age groups was higher for recent than remote items on Day 1, b = 4.5, t(238) = 3.78, p < .001, and on Day 14, b = 20.08, t(238) = 15.78, p < .001. Moreover, this difference between recent and remote items was more pronounced on Day 14 compared to Day 1, b = -15.5, t(238) = -8.88, p < .001. In sum, the results showed that children had overall worse memory retention rates compared young adults, indicating less robust memory consolidation in children. Memory retention immediately after learning were higher compared to overnight and 2-week-old memory retention. This pattern was similarly more pronounced on Day 1 compared to Day 14 for both age groups.

fMRI Results

Mean activation for remote > recent memory in ROIs

Earlier version: In the following section, the results of the univariate analysis of the selected ROIs are summarized, with a full statistical report on LME-model in Table 2 and model-based post hoc Sidak-corrected comparisons in Table S6. Results for the whole-brain analyses are available in Tables S3-5.

Revised version: To investigate the change in the neural activation for correctly recalled memories from short to long delay, we analysed the difference in neural activation for the contrast remote > recent across age groups and sessions during the object presentation time window. We controlled for sex, handedness, general intelligence score, and mean reaction time. In the following section, the results of the univariate analysis of the selected ROIs based on the object presentation tim window (Fig. 1A “Retrieval fMRI) are summarized, with a full statistical report on LME-model in Supplementary Table S6. Results for the whole-brain analyses are available in Supplementary Tables S3-5. All main and interaction effects are adjusted for multiple comparisons with False Discovery Rate (FDR). All post hoc tests were Sidak-corrected.

Added in the revised version: Our results showed that for the anterior and posterior HC (Fig. 3A) as well as for the anterior PHG (Fig. 3B), the mean signal difference for the contrast of remote > recent remained similar across age groups and across sessions (all p > .430 FDR-adjusted), indicating similarly elevated mean blood oxygen level-dependent (BOLD) signal intensity for recent and remote memories across time in both age groups. An additional analysis conducted for recent and remote neural activation measures (for more detailed results refer to Fig. S2 and Table S7) revealed that all activations measures in both age groups we significantly higher than zero (all p < .028FDR-adjusted) other than for recent Day1 posterior hippocampus in children (p = .14FDR-adjusted).

Added in the revised version: Mean Signal Differences Between Correct Remote and Recent Memories.

Removed in the revised version: Our results showed that for the anterior and posterior HC (Fig. 3A) as well as for the anterior PHG (Fig. 3B), difference in neural activation for the contrast of remote > recent remained similar across age groups and across sessions (all p > .450). However, we found a significant increase in remote > recent activation from Day 1 to Day 14 in the posterior PHG (Fig. 3B) in young adults (p = .006), but not in children (p = .66). This increase was significantly higher in adults than in children (p = .009). In contrast, we found that the activation difference in remote > recent in the mPFC (Fig. 3C) was higher in children than in young adults (p = .026), and it remained stable across time (p = .450), indicating a stable involvement of the mPFC in the retrieval of remote memories. Regarding the vlPFC (Fig. 3D), the results showed a significant increase in remote > recent activation from Day 1 to Day 14 (p = .005), while overall activation was higher in young adults in comparison to children (p <.001). In the cerebellum (Fig. 3E), the results revealed an increase from Day 1 to Day 14 in young adults (p = .005), as well as an overall higher remote > recent difference in young adults than in children (p = .024). Moreover, we found a significant decrease in the recruitment of the RSC (Fig. 3F) for remote > recent difference in children from Day 1 to Day 14 (p = .007). This decrease was also more pronounced in children than in adults (p = .012), indicating that over time children recruited the RSC less for remote memories, while the recruitment in young adults remained stable. We also found a significant decrease in the recruitment of the precuneus (Fig. 3G) for remote > recent difference from Day 1 to Day 14 (p = .011). Overall, activation difference remote > recent was lower in children than in young adults (p = .012). Finally, remote > recent activation in the LOC (Fig. 3H) increased from Day 1 to Day 14 in young adults (p <.001), but not in children (p = .42). This effect was significantly higher in young adults than in children (p = .035), indicating that over time, LOC activation increases for successful retrieval of old detail-rich memories in young adults.

Removed in the revised version: Model-based results for ROIs.

Earlier version: Mean signal difference for remote > recent memories on Day 1 and Day 14 in (A) anterior and posterior hippocampus; anterior and posterior parahippocampal gyrus; (C) medial prefrontal cortex; (D) ventrolateral prefrontal cortex; E) cerebellum; (F) retrosplenial cortex; (G) precuneus; (H) lateral occipital cortex. Note: Error bars indicate standard error based on the underlying LME-model. *p < .05; **p < .01; ***p < .001(significant difference); non-significant differences were not specifically highlighted.

Revised version: The figure presents mean signal difference for remote > recent memories on Day 1 and Day 14 in children and adults during the object presentation time window in (A) anterior and posterior hippocampus; (B) anterior an posterior parahippocampal gyrus; (C) cerebellum; (D) medial prefrontal cortex; (E)ventrolateral prefrontal cortex; cerebellum; (F) precuneus; (G) retrosplenial cortex; (H) lateral occipital cortex. Note: Bars represent the average signal difference. The colour indicated the age groups: purple for children and khaki yellow for young adults. Solid-lined bars represent data from Day 1, while dashed-lined bars depict data from Day 14. Across all panels, mean of individual subject data are shown with transparent points. The connecting faint lines reflect within-subject differences across sessions. Error bars indicate standard error of the mean. *p < .05; **p < .01; ***p < .001(significant difference); non-significant differences were not specifically highlighted. Significance main and interaction effects are highlighted by the corresponding asterisks. All main and interactions p-values were FDR-adjusted for multiple comparisons.

Added in the revised version: For the posterior PHG (Fig. 3B), we observed a significant Session x Group interaction, F(1,83) = 9.54, p = .020FDR-adjusted, w2 = .09, indicating more pronounced increase in remote > recent mean signal difference from Day 1 to Day 14 in young adults compared to children, b = .11, t(83) = 3.09, p = .003. Similarly, also in the cerebellum (Fig. 3C) a significant Session x Group interaction, F(1,161) = 7.68, p = .020FDR-adjusted, w2 = .04, indicated stronger increase in remote > recent mean signal difference from Day 1 to Day 14 in young adults compared to children, b = .09, t(160) = 2.77, p = .006.

Added in the revised version: For the mPFC (Fig. 3D), there was a significant main effect of Group, F(1,86) = 7.61, p = .023FDR-adjusted, w2 = .07, denoting lower remote > recent mean signal difference in young adults compared to children, b = -.10, t(86) = −2.76, p = .007. In the vlPFC (Fig. 3E), a significant main effect of Group, F(1,82) = 31.35, p = <.001FDR-adjusted, w2 = .13, highlighted lower remote > recent mean signal difference in children compared to young adults, b = -.125, t(108) = −3.91, p < .001. In addition, in the vlPFC (Fig. 3E), we observed a significant main effect of Session, F(1,99) = 10.68, p = .005FDR-adjusted, w2 = .09, pointing out that remote > recent mean signal difference was higher on Day 14 compared to Day 1, b = .08, t(99) = 3.27, p = .001.

Added in the revised version: In the precuneus (Fig. 3F), a significant main effects were observed for both Group, F(1,161) = 5.09, p = .027FDR-adjusted, w2 = .02, and Session, F(1,161) = 6.50, p = .036FDR-adjusted, w2 = .03. There was a lower remote > recent mean signal difference in adults compared to children, b = -.05, t(160) = −2.26, p = .025, and for Day 14 compared to Day 1, b = -.05, t(160) = - 2.55, p = .012. For the RSC (Fig. 3G), a significant Session x Group interaction, F(1,161) = 8.56, p = .020FDR-adjusted, w2 = .04, showed a greater decrease in remote > recent mean signal difference from Day 1 to Day 14 in children than in young adults, b = -.10, t(160) = −2.93, p = .004. In the LOC (Fig. 3H), a significant main effect of Group, F(1,82) = 9.12, p = .015FDR-adjusted, w2 = .09, indicated a higher remote > recent mean signal difference in young adults compared to children, b = .07, t(82) = 3.02, p = .003. Additionally, a significant main effect of Session, F(1,97) = 16.76, p = <.001FDR-adjusted, w2 = .14, showed an increase in remote > recent mean signal difference on Day 14 compared Day 1 across age groups, b = .07, t(97) = 4.10, p = <.001. Furthermore, a significant Session x Group interaction, F(1,81) = 6.42, p = .032FDR-adjusted, w2 = .06, demonstrated higher increase in remote > recent mean signal difference from Day 1 to Day 14 in adults compared to children, b = .09, t(81) = 2.53, p = .013.

Earlier version: In summary, our results showed increased upregulation of neural activity from Day 1 to Day 14 for remembering remote > recent items in the posterior PHG, LOC and cerebellum in young adults only. Children showed time-invariant remote > recent activation in the mPFC, coupled with a decrease in contrast difference in the precuneus and the RSC from Day 1 to Day 14. Both young adults and children showed increased upregulation of neural activity for remote > recent difference in vlPFC over time. No changes were observed in the anterior and posterior HC and the anterior PHG for either group.

Revised version: In summary, our findings revealed distinct consolidation-related neural upregulation for remote memory between children and adults. From Day 1 to Day 14, adults showed higher increase in remote > recent signal difference for remembered items in the posterior PHG, LOC, and cerebellum than children. Adults showed overall higher remote > recent difference in the vlPFC than children, while children showed overall higher remote > recent difference in the mPFC than adults. Furthermore, we observed a constant activation of anterior and posterior HC and anterior PHG in memory retrieval across age groups irrespective of memory type or delay.

Added in the revised version: Neural-behavioural Correlation

Added in the revised version: We further investigated whether neural upregulation (i.e., remote > recent univariate signal difference) is related to memory performance. Specifically, considering all ROIs simultaneously and differential directionality of remote > recent signal differences, we investigated whether any specific profile of ROI constellation of neural upregulation is related to variations in memory performance. For this purpose, we employed the partial least square correlation analysis (PLSC; Abdi, 2010; Abdi & Williams, 2013). With regard to the interconnectedness of the predefined ROIs, the PLSC is a well-suited method to address multivariate associations between neural measures and memory measures. Consequently, latent variables that represent differential profiles of ROI’s neural upregulations with robust relation with either short- or long-delay variations in memory performance were extracted (for more detailed description of the PLSC method, refer to Method section). In addition, we derived for each subject a value that denotes a within-person robust expression of either short- or long-delay brain profile.

Added in the revised version: For each delay, the permutation test of significance resulted in a single latent variable that reliably and optimally represents across age groups (i) the associations of short delay ROI neural upregulations with variations in short-delay memory accuracy (Fig. 4A; r = .536, p = .0026); and (ii) the associations of long delay ROI neural upregulations with variations in long-delay memory accuracy (Fig. 4C; r = .542, p = .0024). With further bootstrapping, we identified Z-scores estimates of robustness (larger/smaller than ± _1.96 (a < 0.05)) of the components within the multivariate brain profiles across all participants. Thus, for short delay, we observed that higher memory accuracy was robustly associated with greater neural upregulations in the anterior PHG (Z-score = 2.161, r = .347) and vlPFC (Z-score = 3.457, r = .640), as well as with lesser neural upregulation in precuneus (Z-score = −2.133, r = -.323) and cerebellum (Z-score = −2.166, r = -.371) across age groups. In contrast, for long delay, we observed that higher memory accuracy was robustly associated with greater neural upregulation in the vlPFC (Z-score = 3.702, r = .492), RSC (Z-score = 4.048, r = .524), and LOC (Z-score = 3.568, r = .455), and with lesser neural upregulation in mPFC (Z-score = −2.958, r = -.394) across age groups. The identified latent variables indicate that substantial amount of variance (short delay: r = .536 and long delay: r = .542) in either short- or long-delay memory performance was accounted by the identified differential functional profiles of brain regions.

Added in the revised version: Multivariate short- and long-delay brain profiles of neural upregulation (remote versus recent neural activation differences) are associated with variations in memory accuracy.

Added in the revised version: A) Short Delay Brain Profile. Latent variables weights or saliences for each ROI build up one latent variable that expresses a composite short-delay brain profile. Stability of salience elements is defined by Z-scores (depicted as red line: a value larger/smaller than ± _1.96 is treated as reliably robust at (a <0.05). B) Association between Short Delay Retention Rate and Short Delay Brain Score. Within-participant short delay brain scores that represents a within-participant robust expression of the defined latent variable’s profile is plotted against short delay memory retention rates defined as percentage of correctly recalled items on Day 1 relative to Day 0. C) Long Delay Brain Profile. Latent variables weights or saliences for each ROI build up one latent variable that expresses a composite long-delay brain profile. D) Association between Long Delay Retention Rate and Long Delay Brain Score. Within-participant long delay brain scores that represents a within-participant robust expression of the defined latent variable’s profile is plotted against long delay memory retention rates defined as percentage of correctly recalled items on Day 14 relative to Day 0. Note: PHGa – anterior parahippocampal gyrus; PHGp – posterior parahippocampal gyrus; HCa – anterior hippocampus; HCp – posterior hippocampus; PC– precuneous; vlPFC – ventrolateral prefrontal cortex; mPFC – medial prefrontal cortex; RSC – retrosplenial cortex; LOC – lateral occipital cortex; CE – cerebellum; r – Spearman’s rank order correlation index.

Added in the revised version: Identified brain profiles across groups suggest shared patterns between neural mean signal differences in differential sets of ROIs and memory accuracy are consistent across children and adults. However, the strength of this relationship may still differ. To investigate this, we examined with linear regression whether brain score (i.e., weights of the latent variable) predict memory retention rates differentially in the two groups. The results revealed that this relationship was similar between both age groups, as highlighted by non-significant Brain Score x Group interactions for both short delay, F = .52 p = .473, w2 = .00, and for long delay, F = 3.67 p = .059, w2 = .03. Based on this, we ran Spearman’s rank-order correlation analyses across both age groups to identify the strength of the relationship. For short delay, we observed that the stronger expression of brain score was moderately associated with higher memory performance (Fig. 4B), r = .456, p < .001FDR-adjusted. Furthermore, for long delay, the results showed that stronger expression of brain score was also moderately associated with higher long-delay memory performance (Fig. 4D), r(76) = .473, p < .001FDR-adjusted.

Added in the revised version: Taken together, differential short- and long-delay brain profiles of neural upregulation were related to variations in memory accuracy. Despite age-related differences in the derived brain scores, higher expression of within-participant brain score was associated with higher memory retention rates in short and long delay similarly in children and young adults.

Earlier version: Representational (dis)similarity results

Revised version: Representational similarity results

Added in the revised version: In addition to distinct univariate neural upregulation for recent and remote memories, children and adults may exhibit differences in neural representations of these memories. Over time, these representations could also undergo consolidation-related transformations. To address this further, we investigated both more differentiated detailed scene-specific and more generic gist-like neural representations in children and adult.

Corrected scene-specific reinstatement

Added in the revised version: To measure how scene-specific reinstatement at retrieval during fixation time window (after short cue by object presentation; see Fig 1A (Retrieval) and Fig. 5A) changes over time as memories decay, we computed a scene-specific reinstatement index for each neural RSM. We hypothesized that neural patterns evoked by reinstatement of a specific scene without any visual input during fixation time window would be similar to neural patterns evoked by actual presentation of the scene during the scene time window. Therefore, the scene time window was used as a template against which the fixation period can be compared to. Participants were explicitly instructed to recall and visualize the scene and location of the object during fixation time window after being cued by the object. Since the locations were contextually bound to the scene and each object had a unique location in each scene, the location of the object was always embedded in the specific scene context.

Removed in the revised version: The following section presents the results of scene-specific neural reinstatement in the selected ROIs (see Fig. 4A for index calculation overview), with a full statistical report on the LME-model in Table 3 and model-based post-hoc Sidak-corrected comparisons in Table S9. First, we combined the scene-specific reinstatement indices for recent items across sessions, as there were no significant differences between sessions in ROIs in children (all p > .68) and adults (p > .14). All scene-specific reinstatement indices for remote and recent items were significantly different from zero (all p < .05FDR-adjusted; Table S7). The only exceptions were for remote items on Day 14 in the PHG, the mPFC, the RSC, and the precuneus for children.

Earlier version: Representational (Dis)similarity Analysis.

Revised version: Representational Similarity Analysis.

Earlier version: (A) Index Computation (Scene). A representational dissimilarity index was computed by assessing the average dissimilarity between fixation and scene time window separately for recent, remote (Day 1), and remote (Day 14) scenes. (B) Scene-specific index computation. A corrected scene-specific reinstatement index was computed by assessing the average dissimilarity in fixation and scene time window within each trial and subtracting it from the average set dissimilarity between the fixation and scene time window across trials. (C) Index Computation (Category). A representational dissimilarity index was computed by assessing the average dissimilarity for fixation time window for within-category and between-category scenes separately for recent, remote (Day 1), and remote (Day 14) scenes. (D) Category-based index computation. A category-based reinstatement index was computed by assessing the average dissimilarity in fixation time window for same-category pairs and subtracting it from the any-other-category pairs. S – scene time window; F – fixation time window; r – Pearson’s correlation index.

Revised version: (A) Index Computation (Scene). A representational similarity index was computed by assessing the averag similarity between fixation and scene time window separately for recent, remote (Day 1), and remote (Day 14) scenes. (B) Scene-Specific Reinstatement. A corrected scene-specific reinstatement index was computed b assessing the average similarity in fixation and scene time window within each trial and subtracting from it the average set similarity between the fixation and scene time window across trials. S – scene time window; F – fixation time window; r – Pearson’s correlation index; Δ z – difference between two Fisher transformed r values. * - Activation patterns.

Added in the revised version: To investigate how successful scene-specific reinstatement changes over time with memory consolidation, all analyses were restricted to correctly remembered items (Fig. 5). For each specific scene, the correlation between neural patterns during fixation “fixation period” and neural patterns when viewing the scene “scene period” was conducted (Fisher-transformed Pearson’s r; Fig. 5B). A set-based reinstatement index was calculated as an average distance between “fixation” and “scene period” for a scene and every other scene within the stimuli set (Deng et al., 2021; Ritchey et al., 2013; Wing et al., 2015). The set-based reinstatement index reflects the baseline level of non-specific neural activation patterns during reinstatement. We then calculated the corrected scene-specific reinstatement index as the difference between set-based and scene-specific Fisher-transformed Pearson’s r (Deng et al., 2021; Ritchey et al., 2013; Wing et al., 2015). Given the temporal proximity of the fixation and scene time window, we refrain from interpreting the absolute values of the observed scene-specific reinstatement index. However, given that the retrieval procedure is the same over time and presumably similarly influenced by the temporal autocorrelations, we focus primarily on the changes in reinstatement index for correctly retrieved memories across immediate, short, and long delays. In other words, the focus in the following analysis lies on the time-related change in the scene-specific reinstatement index.

Added in the revised version: First, we combined the scene-specific reinstatement indices for recent items across sessions, as there were no significant differences between sessions in ROIs in children (all p > .999) and adults (p > .999). To investigate time-dependent change in scene-specific reinstatement in children and young adults in the predefined ROIs, we conducted a LMER model, with delay (recent, remote short and remote long delays), group (children and young adults) for each ROI, controlling for ROI BOLD activation (Varga et al., 2023) during corresponding sessions. All main and interaction effects were FDR-adjusted and all post hoc tests were Sidak-corrected for multiple comparisons.

Added in the revised version: Generally, in all predefined ROIs, we observed a significant main effect of Session (all p < .001FDR-adjusted) in all ROIs and a significant effect of Group in all ROI (all p <.004 FDR-adjusted), except for the LOC, F(1,100) = 1.23, p = .271, ω2 = .002 (Fig. 6). The pattern of time-related decline was similar across age groups, as indicated by not significant Session x Group interactions in all ROIs (all p > .159). There was no significant effect of BOLD activation (all p > .136). The full statistical report on the LME-model is in Supplementary Material in Table S8. A more detailed overview of the observed main effects and their Sidak-corrected post-hoc tests are summarized in the Table 2.

Removed in the revised version: We observed significant group main effects for all ROIs (all p < .003) except for the vlPFC (p = .892). Namely, our results showed higher overall scene-specific reinstatement in young adults compared to children in the HC (p <.001; Fig. 5A), PHG (p = .003; Fig. 5B), cerebellum (p < .001; Fig. 5C), LOC (p <. 001; Fig. 5D), mPFC (p < .001; Fig. 5E), RSC (p <. 001; Fig. 5G), and precuneus (p < .001; Fig. 5H), but not in vlPFC (p = .89; Fig. 5F). These results indicate that scene-specific reinstatement was generally more attenuated in children than in young adults. Concerning significant session effect for all ROIs (all p < .001), the results showed a significant decrease from recent to short delay in all ROIs (all p <.001), while a further significant decrease from short to long delay was observed only in the vlPFC (p = .039) and the RSC (p = .012). We did not observe any group x session interactions (all p > .077). These results indicate that the main decrease in scene-specific neural reinstatement for successfully consolidated memories occurs after a short overnight delay.

Corrected scene-specific neural reinstatement.

Added in the revised version: Scene-specific neural reinstatement defined as the difference between Fisher-transformed scene-specific and set-specific representational similarity. (A) Hippocampus Anterior; (B) Hippocampus Posterior; (C) Parahippocampal Gyrus Anterior; (D) Parahippocampal Gyrus Posterior; (E) Cerebellum; (F) Medial Prefrontal Cortex; (G) Ventrolateral Prefrontal Cortex; (H) Precuneus; (I) Retrosplenial Cortex; (J) Lateral Occipital Cortex. *p < .05; **p < .01; ***p < .001(significant difference). Error bars indicate standard error. Δ z – difference between two Fisher transformed r values.

Removed in the revised version: All FDR-adjusted scene-specific reinstatement indices that were significantly different from zero were marked with green asterisk (Table S6). (A) Hippocampus; (B) Parahippocampal Gyrus; (C) Cerebellum; (D) Lateral Occipital Cortex; (E) Medial Prefrontal Cortex; (F) Ventrolateral Prefrontal Cortex; (G) Retrosplenial Cortex; (H) Precuneus.

Earlier version: Taken together, we observed more attenuated scene-specific neural reinstatement in children compared to young adults. Scene-specific reinstatement declined significantly for overnight-old memories compared to immediate memories in all ROIs and remained further stable after a 2-week-period for all ROIs, except for the vlPFC and the RSC.

Revised version: Taken together, we observed more attenuated scene-specific neural reinstatement in children compared to young adults. Scene-specific reinstatement declined significantly for overnight old memories compared to immediate memories declined further after a 2-week-period for all ROIs. These results indicate that the main decrease in scene-specific neural reinstatement for successfully consolidated memories occurs already after a short overnight delay and proceeds further after longer fortnight delay.

Added in the revised version: Gist-like neural reinstatement

Added in the revised version: To assess the quality of reinstatement of the scenes belonging to the same category (e.g., field, forest, etc.) during the fixation time window following the object cueing (see Fig. 1A (Retrieval) and Fig. 7), we computed the gist-like reinstatement index. The distribution of within-category items across runs was similar and balanced. Additionally, their presentations within runs were randomised without close temporal proximity. First, a within-category similarity indices were computed based on fixation time window of correctly remembered items belonging to the same category (i.e., field, water, housing, forest, infrastructure, indoor, farming), excluding the similarity computation for the fixation time windows with itself. A between-category similarity indices were computed based on fixation time window of correctly remembered items belonging to different categories. A gist-like reinstatement index was computed by subtracting between-categories from within-categories Fischer-transformed distances ([within categoryrecent r – between categoryrecent r] and [within categoryremote r – between categoryremote r] for each session, Fig. 7) . Therefore, the gist-like reinstatement gives us a measure of the preactivation of the whole category of scenes (i.e., forests).

Added in the revised version: Representational Similarity Analysis.

Added in the revised version: (A) Index Computation (Gist). A representational similarity index was computed by assessing the averag similarity for fixation time window for within-category and between-category scenes separately for recent, remote (Day 1), and remote (Day 14) scenes. The diagonal (similarity of fixation time window with itself) was exclude from the analysis. (B) Gist-like Reinstatement. A gist-like reinstatement index was computed by assessing the average similarity in fixation time window for the same-category pairs and subtracting from it the any-other-category pairs. S – scene time window; F – fixation time window; r – Pearson’s correlation index. Δ z – difference between two Fisher transformed r values.

Added in the revised version: The non-zero values in this index reflect gist-like reinstatement, as the similarity distance would be higher for pairs of trials within the same category, indicating more generic reinstatement (e.g., during reinstatement of scenes belonging to a category “forest”, participants may tend to recall a generic image of some forest without any specific details). In other words, the reinstatement of a more generic, gist-like image of a forest across multiple trials should yield more similar neural activation patterns. Not significant gist-like reinstatement would indicate that even within the same category, reinstatement of specific scenes is sufficiently differential and rich in details, rendering them dissimilar (e.g., participants may tend to recall detailed image of forests: fall forest with yellow trees, dark pike-tree forest, light summer forest with young birch trees, etc.).

Added in the revised version: First, we aggregated the gist-like reinstatement indices for recent items on Day 1 and Day 14, as there were no significant differences between sessions in ROIs in children (all p > .95) and adults (p > .99). Then we applied a one-sample permutation t-test to test for significance of all gist-like indices against zero in each ROI (for full overview see Table S10, Figure S4). FDR-corrected values revealed that young adults did not show any category-based reinstatement (all p > .127), while significant gist-like reinstatement was observed in children in the mPFC, Precuneus, and anterior HC (all p < .042). Following this, we conducted a final LME model, separately for each ROI that showed significant gist-like reinstatement, with Subject as the random factor and Delay (recent, remote Day 1, remote Day 14) and Group (children, young adults) as fixed factors, controlling for the BOLD mean activation in each ROI during corresponding sessions.

Added in the revised version: Second, we investigated the time-dependent change in gist-like reinstatement in ROIs that showed significant gist-like reinstatement. We observed a significant main effect of Group in the mPFC, F(1,75) = 6.77FDR-adjusted, p = .011, ω2 = .03 (Fig. 8B), indicating significantly higher gist-like reinstatement in the mPFC in children compared to young adult, b = .02, t(83) = 2.52, p = .013, 95% CI [.004 – .036]. Neither anterior HC nor precuneous showed any significant main or interaction effects (all p > .111; Fig. 8A and 8C; detailed overview in Table S11). Taken together, only the child group showed gist-like reinstatement in the medial-temporal, medial prefrontal, and parietal brain regions. We observed a significantly higher overall gist-like reinstatement in medial prefrontal cortex region in children compared to young adults, indicating a higher level of gist-like representations in children.

Added in the revised version: Gist-like Reinstatement.

Removed in the revised version: Category-based neural reinstatement

Removed in the revised version: In the following section, the results of the category-based neural reinstatement analyses with the selected ROIs are summarized (see Fig. 4B for index calculation overview),, with a complete statistical report on the LME-model in Table 4 and model-based post-hoc Sidak-corrected comparisons in Table S10. First, we combined the category-based reinstatement indices for recent fixations across sessions, as there were no significant differences between sessions in ROIs in children (all p > .64) and adults (p > .87). We tested whether category-based reinstatement significantly differed from zero in all ROIs for recent (immediate after learning) and remote items (Day 1, Day14) for both age groups (for full overview see Table S8). FDR-corrected values revealed that young adults did not show any category-based reinstatement (all p > .51), while this was true for children only in the LOC (all p > . 15).

Removed in the revised version: We observed higher overall category-based reinstatement in children compared to young adults in the HC (p <.001; Fig. 6A), PHG (p < .001; Fig. 6B), cerebellum (p < .001; Fig. 6C), mPFC (p < .001; Fig. 6E), vlPFC (p <. 001; Fig. 6F), RSC (p <. 001; Fig. 6G), and precuneus (p < .001; Fig. 6H), but not in LOC (p = .052; Fig. 5D). These results indicate that overall category-based reinstatement was generally higher in children compared to young adults. We did not observe any session effects (all p > .136). Moreover, category-based reinstatement for remote items from Day1 to Day 14 was significantly higher in children, compared to young adults, in the HC (p = .002) and in the PHG (p = .014).

Removed in the revised version: Category-based reinstatement.

Earlier version: Category-based reinstatement depicts the difference between pattern dissimilarity for within-category items and between-category items during fixation period before the actual scenes were shown. Higher values mean higher category-based reinstatement. The index was tested for significance against zero and all results were FDR corrected for multiple comparisons. Significant reinstatement of category-based information is depicted by a green asterisk (*) (A) Hippocampus; (B) Parahippocampal Gyrus; (C) Cerebellum; (D) Lateral Occipital Cortex; (E) Medial Prefrontal Cortex; (F) Ventrolateral Prefrontal Cortex; (G) Retrosplenial Cortex; (H) Precuneus; *p < .05; **p < .01; ***p < .001(significant difference); non-significant difference was not specifically highlighted. Error bars indicate standard error.

Revised version: Gist-like reinstatement is reflected by the difference in Fisher’s z (Δ z) between within-category and between-category representational similarity during fixation time window, where participants were instructed to reinstate the scene associated with the learned object before the actual scenes were shown. Higher values mean higher gist-lik reinstatement. The index was tested for significance against zero and all results were FDR corrected for multiple comparisons. Significant reinstatement of gist-like information is highlighted by a green rectangle (A) Hippocampus Anterior; (B) Medial Prefrontal Cortex; (C) Precuneus; *p < .05; **p < .01; ***p < .001(significant difference); non-significant difference was not specifically highlighted. Error bars indicate standard error.

Removed in the revised version: Taken together, only the child group showed category-based reinstatement in the medial-temporal, the cerebellar, the prefrontal, and the parietal brain regions. We observed a significant increase in category-based reinstatement in medial-temporal brain regions over time in children, indicating a higher level of gist-like representations in long delay memories after a 2-week-period.

Neural-behavioural Correlations

Added in the revised version: Further, we also explored whether over time, short- and long-delay scene-specific and gist-like reinstatement is beneficial or detrimental for memory performance by correlating the indices with memory retention rates. We derived, with a PLSC analysis, latent brain pattern across implicated ROIs for reinstatement indices that share the most variance with either short-delay or long-delay variations in memory accuracy.

Added in the revised version: For gist-like reinstatement, we included only those ROIs that showed significant reinstatement (i.e., only in children; mPFC, anterior HC and PC for short delay; mPFC for long delay). For the scene-specific reinstatement also all predefined ROIs in both age groups were included. Finally, we examined how scene-specific and gist-like reinstatement brain profiles are related to memory performance for both children and young adults, correlating these values with memory accuracy for respective delays.

Removed in the revised version: Finally, we examined whether item-specific and category-based neural reinstatement is beneficial to memory performance for both children and young adults, correlating reinstatement indices with memory retention rate for Day 14. First, we combined the indices across ROIS, as there were no significant differences between ROIs in relation to reinstatement index x retention rate correlation (all p > .44). The results revealed that a higher item-specific reinstatement index (Fig.7A) was positively related to memory retention rates in children, r = .39, t = 2.44, p = .02, (p = .044FDR adjusted), and in young adults, r = .35, t = 2.10, p = .044, (p = .044FDR adjusted). Furthermore, higher category-based reactivation index (Fig. 7B) was negatively related to memory retention rates in children, r = -.53, t = -3.63, p = .0009, (p = .0018FDR adjusted), but not in young adults, p = .22. Taken together, more differentiated detail-rich neural reinstatement was related to better recollection of consolidated memories in both children and young adults. On the other hand, more gist-like neural reinstatement was related to worse recollection of consolidated memories in children but not young adults.

Earlier version: Neural-behavioural Correlations.

Revised version: Neural-behavioural correlations (scene-specific reinstatement)

Added in the revised version: First, for short delay, the permutation test of significance resulted in a single latent variable that robustly represents the association of scene-specific reinstatement brain profile (Fig. 9A) and memory accuracy across both age groups (Fig. 9B, r = .339, p = .0017). With further bootstrapping we identified Z-scores estimates of robustness (larger/smaller than ± _1.96 (a < 0.05)) of the components within the multivariate brain profile. Thus, for short delay, we observed that higher memory accuracy was robustly associated with greater scene-specific reinstatement in the anterior PHG (Z-score = 2.885, r = .371), posterior PHG (Z-score = 2.597, r = .342), anterior HC (Z-score = 3.126, r = .399), posterior HC (Z-score = 2.844, r = .375), vlPFC (Z-score = 2.434, r = .317), mPFC (Z-score = 2.753, r = .333), and LOC (Z-score = 2.176, r = .298) across age groups.

Added in the revised version: Multivariate short- and long-delay brain profiles of scene-specific reinstatement are associated with variations in memory accuracy.

Added in the revised version: A) Short Delay Brain Profile. Latent variables weights or saliences for each ROI build up one latent variable that expresses a composite short-delay scene-specific reinstatement brain profile. Stability of salience elements is defined by Z-scores (depicted as red line: a value larger/smaller than ± _1.96 is treated as reliably robust at (a <0.05). B) Association between Short Delay Retention Rate and Short Delay Scene-Specific Reinstatement Brain Score. Within-participant short delay scene-specific reinstatement brain scores that represents a within-participant robust expression of the defined latent variable’s profile is plotted against short dela memory retention rates defined as percentage of correctly recalled items on Day 1 relative to Day 0. C) Long Delay Brain Profile. Latent variables weights or saliences for each ROI build up one latent variable that expresses a composite long-delay scene-specific reinstatement brain profile. Stability of salience elements is defined by Z-scores (depicted as red line: a value larger/smaller than ± _1.96 is treated as reliably robust at (a <0.05). B) Associatio between Long Delay Retention Rate and Long Delay Scene-Specific Reinstatement Brain Score. Within-participant long delay scene-specific reinstatement brain scores that represents a within-participant robust expression of the defined latent variable’s profile is plotted against long delay memory retention rates defined as percentage of correctly recalled items on Day 14 relative to Day 0. Note: PHGa – anterior parahippocampal gyrus; PHGp – posterior parahippocampal gyrus; HCa – anterior hippocampus; HCp – posterior hippocampus; PC– precuneous; vlPFC – ventrolateral prefrontal cortex; mPFC – medial prefrontal cortex; RSC – retrosplenial cortex; LOC – lateral occipital cortex; CE – cerebellum; r – Spearman’s rank order correlation index.

Added in the revised version: Second, for long delay, the permutation test of significance resulted in a single latent variable that robustly represents the association of scene-specific reinstatement brain profile (Fig. 9C) and memory accuracy across both age groups (Fig. 9D, r = .455, p = <.001). Further, for long delay, we observed that higher memory accuracy was robustly associated with greater scene-specific reinstatement in the anterior PHG (Z-score = 6.213, r = .414), posterior PHG (Z-score = 4.810, r = .334), anterior HC (Z-score = 5.353, r = .389), posterior HC (Z-score = 4.707, r = .354), precuneous (Z-score = 3.404, r = .281), vlPFC (Z-score = 3.291, r = .266), RSC (Z-score = 3.72, r = .293), LOC (Z-score = 3.288, r = .282), and cerebellum (Z-score = 3.842, r = .308) across age groups.

Added in the revised version: Further, the linear regression analysis revealed similar relationship between identified brain profiles and memory accuracy between children and adult as indicated by non-significant Scene-Specific Reinstatement Brain Score x Group interactions for both short delay, F = 2.61 p = .110, w2 = .02, and for long delay, F = .43 p = .836, w2 = .00. Based on this, we ran Spearman’s rank-order correlation analyses across both age groups to identify the strength of the relationship. For short delay, we observed that the stronger expression of scene-specific reinstatement brain score was moderately associated with higher short-delay memory retention rate (Fig. 8B), r = .413, p < .001FDR-adjusted. Furthermore, for long delay, the results showed that stronger expression of scene-specific reinstatement brain score was also moderately associated with higher long-delay memory retention rates (Fig. 8D), r = .419, p < .001FDR-adjusted. These significant correlations underscore the importance of scene-specific reinstatement in positively contributing to memory performance for detailed associative information both in children and adult. The lack of a significant difference between children and adults suggests that the fundamental relationship between scene-specific reinstatement and memory might also remain consistent across age groups.

Added in the revised version: Neural-behavioural correlations (gist-like reinstatement)

Added in the revised version: First, for short delay, the permutation test of significance resulted in a single latent variable that robustly represents the association gist-like reinstatement brain profile (Fig. 10A) and memory accuracy in children (Fig. 10A, r = .379, p = .024). For short delay, we observed that higher memory accuracy was robustly negatively associated with greater gist-like reinstatement in the anterior HC (Z-score = −1.985, r = -.681), and mPFC (Z-score = −2.189, r = - .681) in children.

Added in the revised version: Multivariate short- and long-delay brain profiles of gist-like reinstatement are associated with variations in memory accuracy.

Added in the revised version: A) Short Delay Brain Profile. Latent variables weights or saliences for each ROI build up one latent variable that expresses a composite short-delay gist-like reinstatement brain profile. Stability of salience elements is defined by Z-scores (depicted as red line: a value larger/smaller than ± _1.96 is treated as reliably robust at (a <0.05). B) Association between Short Delay Retention Rate and Short Delay Gist-Like Reinstatement Brai Score. Within-participant short delay gist-like reinstatement brain scores that represents a within-participant robust expression of the defined latent variable’s profile is plotted against short delay memory retention rates defined as percentage of correctly recalled items on Day 1 relative to Day 0. C) Long Delay Brain Profile. Latent variables weights or saliences for each ROI build up one latent variable that expresses a composite long-delay gist-lik reinstatement brain profile. Stability of salience elements is defined by Z-scores (depicted as red line: a value larger/smaller than ± _1.96 is treated as reliably robust at (a <0.05). B) Association between Long Delay Retention Rate and Long Delay Gist-Like Reinstatement Brain Score. Within-participant long delay gist-like reinstatement brain scores that represents a within-participant robust expression of the defined latent variable’s profile is plotted against long delay memory retention rates defined as percentage of correctly recalled items on Day 14 relative to Day 0. Note: HCa – anterior hippocampus; PC– precuneous; mPFC – medial prefrontal cortex; r – Spearman’s rank order correlation index.

Added in the revised version: Second, for long delay, the permutation test of significance resulted in a single latent variable that robustly represents the association of scene-specific reinstatement brain profile (Fig. 10C) and memory accuracy across both age groups (Fig. 10D, r = .372, p = .015). Further, for long delay, we observed that higher memory accuracy was robustly associated with lower gist-like reinstatement in the mPFC (Z-score = −3.354, r = .371) in children.

Added in the revised version: Based on this, we ran Spearman’s rank-order correlation analyses to identify the strength of these relationships in child group. For short delay, we observed a trend-level negative association between stronger expression of gist-like reinstatement brain score and memory performance (Fig. 8B), r = .266, p = .08FDR-adjusted. Furthermore, for long delay, the results showed that stronger expression of gist-like reinstatement brain score was moderately associated with higher long-delay memory retention rates (Fig. 8D), r = .390, p = .02FDR-adjusted. The significant correlation observed in children underscores the importance of gist-like reinstatement in being detrimental to memory performance for detailed associative information in children in long delay.

Added in the revised version: Taken together, more differentiated detail-rich neural reinstatement was related to better memory retrieval in both children and young adults. On the other hand, uniquely in children, more gist-like neural reinstatement was related to worse memory retrieval.

Removed in the revised version: Reinstatement indices were averaged across ROIs with significant reinstatement index. (A) Item-specific reinstatement indices were significantly positively related to long-delay memory retention rates in children (in purple) and young adults (in yellow). (B) Category-based reinstatement indices were significantly negatively related to long-delay memory retention rates in children (in purple), but not in young adults (in yellow). R = correlation coefficient, p = p-value. All p-values were FDR-adjusted for multiple comparisons.

Discussion

Earlier version: In the present study, we investigated system-level memory consolidation of object-location associations after learning with immediate delay, after one night of sleep as short delay and after two weeks as long delay. We tracked changes in neural activation level and multivariate reinstatement patterns over time, comparing 5-to-7-year-old children and young adults. Our main findings are as follows: (i) Children showed overall greater decline in memory retention compared to young adults. (ii) In terms of activation level, both age groups showed similar activation in the HC and the anterior PHG, increase in the vlPFC, and decrease in the precuneus over time. Young adults exhibited higher increase in neural activation in the posterior PHG, the cerebellum, and the LOC over time, and time-invariant higher activation in the vlPFC compared to children. In contrast, children showed higher decrease in the RSC activation over time, and time-invariant higher activation in the mPFC compared to young adults. (iii) Using RSA, we found that children exhibited attenuated scene-specific neural reinstatement compared to young adults. The decline in scene-specific neural reinstatement occurred mainly overnight, while over longer delay it remained stable in all ROIs and declined further only in the vlPFC and the RSC. We observed that category-based or less differentiated gist-like reinstatement was present only in children in medial-temporal, cerebellar, parietal, and prefrontal brain regions. This pattern significantly increased over time in the HC and PHG only in children. Importantly, higher item-specific reinstatement was related to better retention rates in children and young adults, while higher category-based reinstatement was related to lower retention rates only in children.

Revised version: In the present study, we investigated system-level memory consolidation of object-location associations after learning with immediate delay, one night of sleep as short delay and after two weeks as long delay. We tracked changes in neural activation and multivariate reinstatement patterns over time, comparing 5-to-7-year-old children and young adults. Our main findings are as follows: (i) Children showed greater decline in memory retention both in short and long delay compared to young adults. (ii) In terms of neural upregulation, reflected as the mean difference between remote > recent neural activation, age groups showed distinct changes over time. Young adults exhibited increase in neural upregulation in the posterior PHG, cerebellum and LOC over time, as well as overall higher neural upregulation in the vlPFC compared to children . In contrast, only children showed decrease in neural upregulation in the RSC over time, and they showed overall higher neural upregulation in the mPFC than adults. Distinct neural upregulation profiles with a specific set of brain regions were related to short and long delay memory accuracy. (iii) Using RSA, we found that differentiated scene-specific reinstatement was more prominent in adults than children and decreased over time in both age group. We observed that more generic gist-like reinstatement was present only in children in anterior hippocampal and medial prefrontal brain regions. Importantly, higher scene-specific reinstatement was related to better retention rates in both children and young adults, while higher gist-like reinstatement was related to lower retention rates only in children.

Added in the revised version: Our study extends previous adult-based findings and, for the first time, demonstrates that the retrieval of consolidated memories in children is accompanied by differential patterns of neural activation of some of the core retrieval brain regions, attenuated neural reinstatement of detailed specific memories, and stronger generic gist-like reinstatement. Our results suggest that adults can utilize their mature neural memory systems and extensive existing knowledge structure to encode and consolidate new complex information with detailed accuracy. In contrast, children utilize their neural resources, which are still undergoing maturation, to build up their sparse knowledge structures. Their memory system may tend to favour encoding and consolidating gist as a more solid building block for their still sketchy knowledge base, sacrificing detailedness. At this developmental stage, focusing on details may not be a priority (Keresztes et al., 2018). We discuss each finding in detail in the following sections.

Removed in the revised version: Our study extends previous adult-based findings and, for the first time, demonstrates that the retrieval of consolidated memories in children is accompanied by decreased neural activation of core retrieval brain regions, more attenuated reinstatement of memory details, and more category-based, gist-like reinstatement. We discuss each finding in detail in the following sections.

Earlier version: Less robust overall memory retention in children compared to young adults

Revised version: Less robust short and long delay memory retention in children compared to young adults

Added in the revised version: Our findings indicate that preschool 5-to-7-year-old children can encode and retain complex associative and highly contextualized information successfully over extended periods after adaptive learning. However, they had overall lower learning and retrieval performance compared to young adults. In addition, these children exhibited more pronounced declines in retention rates over both short and long delays decrease for correctly learned information, suggesting less robust memory consolidation compared to young adults.

Added in the revised version: Concerning learning, overall children needed more cycles to memorize object-scene associations and showed lower learning performance after initial strategic encoding compared to young adults. Although we did not expect children to show similar learning rates to adults due to the complex and associative nature of the task (Pressley et al., 1981), we aimed to maximize children’s learning capacities through adaptive learning. Therefore, attention allocation and motivation during encoding and learning were controlled for by the constant presence of the experimenter and feedback questionnaires. Moreover, all participants underwent training to create elaborative memories that help to support retrieval.

Added in the revised version: Overall, our findings on learning suggests that children were less adept at utilizing strategic control over encoding by creating and maintaining stories to aid their retrieval as successfully as adults. This is consistent with previous literature, showing continuous improvement in children’s ability to use elaborative strategies between ages 4 and 8 (Bjorkund et al., 2009; Crowley & Siegler, 1999; Pressley, 1982). Additionally, children at this age may experience difficulties in controlling (Ruggeri et al., 2019) and effectively using their learning strategies over time (Brod, 2021; Shing et al., 2010). Observed lower learning rates may also be attributed to less efficient binding processes in children compared to young adults (Shing et al., 2010; Sluzenski et al., 2006). Although we included only stimuli from the primary school curriculum to reduce age differences in knowledge availability, ongoing maturation of the memory brain network in 5-to-7-year-old children may have attenuated their benefit from pre-existing knowledge and memory aid through strategic elaboration (Ghetti & Bunge, 2012; Lenroot & Giedd, 2006; Nishimura et al., 2015; Ofen, 2012; Shing et al., 2008). Despite these challenges, 5-to-7-year-old children were capable of learning complex associative information to a considerable extent, which aligns with their ability to gradually accumulate world knowledge (Bauer, 2021; Brod & Shing, 2022; Wagner, 2010).

Earlier version: Children showed overall less robust mnemonic performance for complex associative information compared to young adults. Specifically, we showed that already during learning, children needed more cycles to memorize object-scene associations and have a lower learning performance, which may be attributed to less efficient binding and strategy use as well as schema-integration processes in children compared to young adults (Shing et al., 2010). Although we included only stimuli from the primary school curriculum to reduce age differences in knowledge availability, ongoing maturation of the memory brain network in 5-to-7-year-old children may have attenuated their benefit from pre-existing knowledge and memory aid through strategic elaboration (Ghetti & Bunge, 2012b; Ofen, 2012; Shing et al., 2008). Furthermore, our results are in line with previous studies that reported worse memory retention for associative information in school age children compared to adults (Østby et al., 2012; Schommartz et al., 2021). On the other hand, our results are not in line with sleep-related beneficial effects on mnemonic performance of 7-to-12-year-old children after one night delay (Peiffer et al., 2020; Wang et al., 2018) that were shown for novel stimuli not related to any prior knowledge (in the sense of arbitrary stimuli). As we opted for well-learned information that should allow for rapid creation of new schemas or integration of new associations into already existing schemas, our findings indicate that the beneficial role of sleep on memory consolidation in children compared to adults may not apply for repeatedly and strategically learned information. Deliberate learning is potentially more advantageous for subsequent memory retention in young adults, as this information may be integrated into pre-existing knowledge structures faster (van Kesteren et al., 2013), with higher strategic control of memories upon retrieval and therefore greater accessibility of consolidated memories (Fandakova et al., 2017; Gaudreau et al., 2001). Taken together, our findings indicate that compared to young adults, 5-to-7-year-old children exhibit less robust memory consolidation for well-learned information, suggesting an overall reduced ability to retain detailed memories in children.

Revised version: Concerning memory consolidation, our results are in line with previous studies that reported worse memory retention for associative information in school age children compared to adults (Østby et al., 2012; Schommartz et al., 2023, 2024). On the other hand, our results are not in line with sleep-related beneficial effects on mnemonic performance of 7-to-12-year-old children after one night delay (Peiffer et al., 2020; Wang et al., 2018) that were shown for novel stimuli not related to any prior knowledge (in the sense of arbitrary stimuli). As we opted for well-learned information that should allow for rapid creation of new schemas or integration of new associations into already existing schemas, our findings indicate that the beneficial role of sleep on memory consolidation in children compared to adults may not apply for repeatedly and strategically learned information. Deliberate learning is potentially more advantageous for subsequent memory retention in young adults, as this information may be integrated into pre-existing knowledge structures faster (van Kesteren et al., 2013), with higher strategic control of memories upon retrieval and therefore greater accessibility of consolidated memories (Fandakova et al., 2017; Gaudreau et al., 2001). Taken together, our findings indicate that compared to young adults, 5-to-7-year-old children exhibit less robust memory consolidation for well-learned information, suggesting an overall reduced ability to retain detailed memories in children.

Added in the revised version: Our findings indicate suggest that lower memory performance in children potentially indicate lower memory strength. Therefore, we conducted exploratory analysis with drift diffusion modelling (Lerche & Voss, 2019; Palada et al., 2016; Ratcliff et al., 2011, 2012; Ratcliff & McKoon, 2008; Zhou et al., 2021), deriving memory strength using as drift rate parameter (see Figure S1 and section S2.1 in Supplementary Materials). Our results demonstrate that children have significantly lower drift rate compared to young adults, indicating slower evidence accumulation and noisier recall. As drift rate closely correlates with memory accuracy (Ratcliff et al., 2011), our findings on the memory strength align with those on memory accuracy during retrieval in both age groups. Our neural findings suggest that differences in functional engagement of the retrieval network and the characteristics of memory representations being created and retained may underlie the observed behavioural differences.

Differential upregulation of remote > recent neural activation over time in children in comparison to young adults

Earlier version: Analyses of remote > recent difference in neural activation over time allowed us to control for the effects of rapid consolidation during repeated learning, examining changes in short- or long-delay neural activation relative to immediate retrieval activation (Brodt et al., 2016a, 2018; Yu et al., 2022). First, we observed increased activation in the vlPFC over time in both age groups, while overall vlPFC activation was higher in young adults. Furthermore, we observed stable activation in the mPFC over time in both age groups, while overall mPFC activation was higher in children. This may indicate that both age groups engaged more strategic control over retrieval processes over time, and schema activation for retrieval irrespective of time delay. This is in line with previous studies that have shown the involvement of mPFC in structured and schema-related retrieval of long-term memories (Takashima et al., 2006; Yamashita et al., 2009), as well as the role of vlPFC in strategic remembering and retrieval of stored memories, executing a cognitive control role over mnemonic interferences during retrieval (Badre & D’Esposito, 2009; Kuhl et al., 2012). Here, we show not only involvement of medial and lateral PFC in remote memories, but that their role increases as consolidation progresses over time in both age groups, supporting the notion of (Milton et al., 2011) on growing prefrontal role during delayed recollection. Over time, cognitive control during memory retrieval may increase as it requires greater effort to recollect elaborative stories to remember the associated spatial context. On average, strategic control over memories may be present but less pronounced in children due to the more protracted developmental trajectories of prefrontal cortex maturation (Ghetti & Bunge, 2012b; Gogtay et al., 2004; Ofen, 2012; Shing et al., 2010b). Children also showed overall higher activation in mPFC during retrieval of short and long delay memories, indicating possible schema-related retrieval that may be mediated by mPFC in children to a greater extent than in young adults, who may rely more on strategic retrieval.

Revised version: Analyses of neural upregulation (i.e., remote > recent difference in neural activation) over time allowed us to control for the effects of rapid consolidation during repeated learning, while examining changes in short- or long-delay neural activation (Brodt et al., 2016b, 2018; Yu et al., 2022). First, we observed increased upregulation in the vlPFC over time in both age groups, with vlPFC upregulation being higher in young adults. Furthermore, we observed stable upregulation in the mPFC over time in both age groups, with the overall mPFC upregulation being higher in children. On the one hand, this may indicate a stronger strategic control over retrieval processes over time in young adult, due to vlPFC’s role in strategic remembering and retrieval of stored memories (Badre & D’Esposito, 2009; Kuhl et al., 2012). Over time, cognitive control during memory retrieval may increase as it requires greater effort to recollect elaborative stories to remember the associated spatial context. Strategic control over memories may be present but less pronounced in children due to the more protracted developmental trajectories of prefrontal cortex maturation (Ghetti & Bunge, 2012c; Gogtay et al., 2004; Ofen, 2012; Shing et al., 2010b). On the other hand, our results indicate a more pronounced schema-related retrieval that may be mediated by mPFC to a greater extent in children than in young adults. This extends previous findings on the involvement of mPFC in structured and schema-related retrieval of long-term memories (Takashima et al., 2006; Yamashita et al., 2009) to a child developmental cohort. Interestingly, higher mPFC upregulation in long delay was negatively related to long delay memory accuracy, suggesting that schema-reliance is detrimental to the retention of detailed associative memories. In addition, it may suggest consolidation-related transformation of memory traces into less differentiated, more generic and gist-like memories(Gilboa & Marlatte, 2017; Gilboa & Moscovitch, 2021).

Earlier version: Second, in other constituents of the recollection network (Ranganath & Ritchey, 2012), we observed increased activation from short to long delay in the posterior PHG and stable activity in the precuneus and the RSC in young adults, while children showed a decrease in posterior brain regions. As young adults showed higher memory retention rates for more detail-rich information, this superior memory may be mediated by higher activation in the posterior PHG involved in contextual associations and scene memory (Aminoff et al., 2013). In children, PHG goes through prolonged maturation (Golarai et al., 2007), and its increased functional maturation is related to long-term scene recollection (Chai, 2010). In addition, higher mnemonic distinctiveness in young adults (i.e., higher retention rates for detailed information) may also be mediated by stable RSC and precuneus activation, as these regions are involved in mnemonic vividness, spatial, and associative memory (Brodt et al., 2016a; Hebscher et al., 2019; Mitchell et al., 2018; Richter et al., 2016; Tambini & D’Esposito, 2020; Vann et al., 2009). Time-related decrease in the posterior brain regions in children is also in line with previous findings (DeMaster & Ghetti, 2013), which showed that the involvement of parietal regions in the recollection of correct memories increased with age in 8-to-11-year-old children. Therefore, the continuing maturation of parietal regions in 5-to-7-year-old children (Sowell et al., 2002) presumably underlied the age-related differences in activation patterns in these regions.

Revised version: Second, in other constituents of the recollection network (Ranganath & Ritchey, 2012), we observed increased in upregulation from short to long delay in the posterior PHG and overall lower upregulation in precuneous (i.e., remote > recent) in young adults, while children showed a corresponding decrease in the RSC. As young adults showed higher memory retention rates for more detail-rich information, this superior memory may be mediated by higher upregulation in the posterior PHG involved in contextual associations and scene memory (Aminoff et al., 2013). In children, PHG goes through prolonged maturation (Golarai et al., 2007), and its increased functional maturation is related to long-term scene recollection (Chai, 2010). In addition, higher mnemonic distinctiveness of more recent memories (i.e., higher retention rates for detailed information) may also be mediated by RSC and precuneous activation profiles, as these regions are involved in mnemonic vividness, spatial, and associative memory as indicated by other findings from immediate delays (Brodt et al., 2016b; Hebscher et al., 2019; Mitchell et al., 2018; Richter et al., 2016; Tambini & D’Esposito, 2020; Vann et al., 2009). Moreover, lower short delay precuneus upregulation and higher long delay RSC upregulation was related to better memory performance. Time-related decrease in the posterior brain regions in children is also in line with previous findings (DeMaster & Ghetti, 2013), which showed that the involvement of parietal regions in the recollection of correct memories increased with age in 8-to-11-year-old children. Therefore, the continuing maturation of parietal regions in 5-to-7-year-old children (Sowell et al., 2002) presumably underlies the age-related differences in consolidation-related upregulation in these regions.

Earlier version: Third, the observed increase in neural activation from short to long delay in the LOC and the cerebellum in young adults is also in line with the previous findings showing that the cerebellum supports rapid cortical storage of memory traces after repeated exposure even after 24 hours (Stroukov et al., 2022), and showed upregulation of neural activation for long-term episodic memory (Andreasen et al., 1999). Previous studies also showed that HC-LOC activation was related to scene-related associative memory consolidation (Tambini et al., 2010), and human object recognition (Grill-Spector et al., 2001). Moreover, the network of angular gyrus and LOC has been shown to enhance the overnight retention of schema-related memories in young adults (van der Linden et al., 2017). Activation increases in the cerebellum and LOC in young adults compared to children indicate that these regions support long-delay memory retention in young adults and undergo development into later childhood.

Revised version: Third, the observed increase in neural upregulation from short to long delay in the LOC and the cerebellum in young adults is also in line with the previous findings showing that the cerebellum supports rapid cortical storage of memory traces after repeated exposure even after 24 hours (Stroukov et al., 2022), and showed upregulation of neural activation for long-term episodic memory (Andreasen et al., 1999). Concerning the LOC, previous studies also showed that HC-LOC activation was related to scene-related associative memory consolidation (Tambini et al., 2010), and human object recognition (Grill-Spector et al., 2001). Moreover, the network of angular gyrus and LOC has been shown to enhance the overnight retention of schema-related memories in young adults (van der Linden et al., 2017). In line with this, we also observed that higher long delay LOC upregulation was related to better memory performance. The more pronounced upregulation from short to long delay in these regions in adults suggests that the cerebellum and LOC support long-delay memory retention and their functional role is underdeveloped in middle childhood.

Earlier version: Finally, our findings on age-group- and delay-invariant activation in the anterior HC and PHG, and posterior HC during the retrieval of detail-rich memories (i.e., the exact location of an object within a scene) are in line with Nadel & Moscovitch (1997),who postulated that the hippocampal formation and related structures remain involved in detail-rich memories upon their retrieval, irrespective of memory age. For example, Du et al. (2019) reported stable hippocampal involvement during retrieval of associative memory across delays of one day, one week and one month in young adults. Furthermore, the absence of age-related differences in HC and anterior PHG involvement are also in line with developmental studies that have reported the relative maturity of the HC in middle childhood (Keresztes et al., 2017; Lee et al., 2014; Shing et al., 2010b), which is concomitant with an improvement in the ability to bind event features together into a coherent representation around the age of six years (Sluzenski et al., 2006). Taken together, the similar engagement of medial-temporal cortex over time in children and adults indicated that the retrieval of well-learned detail-rich memories is mediated by these brain structures already in middle childhood.

Revised version: Finally, our findings on age-group and delay-invariant activation in the anterior HC and PHG, and posterior HC during the retrieval of detail-rich memories (i.e., the exact location of an object within a scene) are in line with Nadel & Moscovitch (1997),who postulated that the hippocampal formation and related structures remain involved in detail-rich memories upon their retrieval, irrespective of memory age. For example, Du et al. (2019) reported stable hippocampal involvement during retrieval of associative memory across delays of one day, one week and one month in young adults. Tanrıverdi et al. (2022) also demonstrated that post-encoding coactivation of hippocampal and cortical brain regions may lead to experience-dependent change in memories, highlighting the importance of hippocampal involvement during consolidation. Furthermore, the absence of age-related differences in HC and anterior PHG involvement are also in line with developmental studies that have reported the relative maturity of the HC in middle childhood (Keresztes et al., 2017; Lee et al., 2014; Shing et al., 2010b), which is concomitant with an improvement in the ability to bind event features together into a coherent representation around the age of six years (Sluzenski et al., 2006). Specifically, our finding on hippocampal engagement being robust in children and adults extends the Multiple Trace Theory to a child developmental cohort (Moscovitch & Gilboa, 2022; Nadel et al., 2000). Taken together, the similar engagement of medial-temporal cortex over time in children and adults indicated that the retrieval of well-learned detail-rich memories is mediated by these brain structures already in middle childhood.

Earlier version: To summarize, we provide novel evidence about changes in neural activation for successfully consolidated memories over short and long delay, relative to immediately learned memories. While we showed adult-like activation in medial-temporal brain regions in children, young adults relied more on prefrontal, occipital, cerebellum, and parietal brain regions over time, compared to more pronounced reliance on medial prefrontal region in children.

Revised version: To summarize, we provide novel evidence about changes in neural upregulation for successfully consolidated memories over short and long delay, relative to immediately learned memories. While children exhibited adult-like stable neural activation for recent and remote memories in medial-temporal brain regions, young adults relied more on prefrontal, occipital, cerebellar, and parietal brain regions over time, compared to more pronounced reliance on medial prefrontal regions in children. Adults show more mature neocortical consolidation-related engagement, resulting in stronger and more durable detailed memories over time while in children immature neocortical engagement may lead to consequent reduction in memory retention of detailed memories.

Earlier version: Reduced scene-specific and enhanced category-based reinstatement in children compared to young adults

Revised version: Reduced scene-specific reinstatement over time in children and young adults

Earlier version: Scene-specific reinstatement was more attenuated in children in comparison to young adults. Concerning the overall time-related decline, we observed that the drop in scene-specific reinstatement occurred mainly overnight with short-delay consolidation, indicating that neural patterns lost mnemonic specificity as memories aged. Moreover, more differentiated reinstatement of scene was related to higher retention rates both in children and young adults.

Revised version: We found that scene-specific reinstatement decreased over time both in children and young adults, aligning with delay-related decrease in memory retention. Additionally, it was overall more attenuated in children compared to young adults. Higher scene-specific neural reinstatement was related to better memory performance in short and long delay in both age groups.

Added in the revised version: Our findings contribute to the memory consolidation literature by demonstrating that scene-specific neural reinstatement observed in neocortical, medial temporal and cerebellar brain regions supports reinstatement of detailed specific contextual memories. This observation is consistent with the Contextual Binding Theory (Yonelinas et al., 2019), which posits that stronger reinstatement of contextual details can enhance memory retention. The similar decay of these processes over time in both children and adults suggests that the basic mechanisms of contextual binding are present early in development. Additionally, in line with the Trace Transformation Theory (Moscovitch & Gilboa, 2022), our findings suggest that reinstatement patterns continuously transform over time. This transformation, observed across all considered memory-related regions, indicates a consistent and systematic consolidation-related reshaping of the unique scene-specific memory representations over time (Chen et al., 2017).

Added in the revised version: Our findings on scene-specific reinstatement align with and add to the previous literature that show reliable reinstatement of unique events. For example, our findings align with the effects observed by Masís-Obando et al. (2022) for the immediate recall of story details in key memory regions. Consistent with Oedekoven et al. (2017), our results show that memory representations for unique events can be reliably detected through scene-specific reinstatement even after extended delays. Furthermore, we build on Guo & Yang (2022) by demonstrating how specific ROI-related profiles of neural reinstatement during retrieval correlate with long-term memory retention. Unlike Oedekoven et al. (2017), who reported no time-related differences in reinstatement effects and used the same video clips for immediate and delayed recall – which could have inadvertently reinforced memory through reactivation – our study employed unique stimulus sets for each retrieval sessions, preventing any reconsolidation of mnemonic representations. This approach revealed a significant attenuation of reinstatement patterns after an overnight delay, which further diminished after two weeks, highlighting the importance of intentional reactivation for maintaining the specificity of neural reinstatement.

Removed in the revised version: Our findings concerning scene-specific reinstatement are in line with the reinstatement effects for immediate recall of story details shown by Masís-Obando et al. (2022) in mPFC, posterior medial cortex, PHG, etc. We show that scene-specific reinstatement can be observed even after longer time delays. Moreover, we extend the findings of Guo & Yang (2022) on neural reinstatement in LOC during encoding and its relatedness to successful memory performance, showing neural reinstatement in LOC for successfully recollected memories over time. Additionally, our findings are in line with the reported reinstatement effects for visual clips shown by Oedekoven et al. (2017) for encoding, immediate recall, and 1-week-delayed recall in precuneus, medial temporal gyrus, occipital gyrus, HC, and PHG. However, in contrast to Oedekoven et al. (2017), who did not observe any time-related differences in the neural reinstatement effect, we showed a major attenuation in reinstatement patterns in all ROIs after an overnight delay, followed by further reinstatement attenuation after a 2-week-period. These discrepancies in findings may be because Oedekoven et al. (2017) used the same set of video clips for immediate and delayed recall, possibly introducing additional retraining or mnemonic reactivation that may have boosted the fidelity of mnemonic reinstatement over time. Contrary to that, we used unique sets of stimuli for each retrieval, avoiding any reconsolidation of the mnemonic representations. Thus, these results show that without intentional reactivation and reconsolidation, the specificity of neural reinstatement fades over time. Additionally, in line with our findings, Xiao et al. (2017) showed item-specific neural reinstatement of scenes in frontoparietal cortex after learning.

Earlier version: In terms of age differences, in contrast to our findings of more attenuated neural reinstatement in children compared to young adults, Fandakova et al. (2019) showed that neural specificity during encoding was similar in 8-to-15-year-old children and adults in RSC, LOC and PHG. However, Fandakova et al. (2019) tested neural specificity with a slightly older cohort of children, suggesting that in late childhood to early adolescence there is already adult-level specificity of neural patterns reinstatement (Ghetti & Bunge, 2012b; Gogtay et al., 2004). Furthermore, our findings on the presence of scene-specific reinstatement in 5-to-7-year-old children, albeit being at a lower level compared to adults, are also in line with previous studies (Benear et al., 2022; Cohen et al., 2019; Golarai et al., 2015) that showed reliable mnemonic reinstatement for visual input (i.e., faces, movie clips) in 5-to-11-year-old children. Moreover, we extend previous finding, showing that successful retrieval of long-term memories is related to more differentiated neural reinstatement both in children and young adults, indicating at similar mechanisms of detail-rich memory consolidation already 5-to-7-year-old children and young adults.

Revised version: Our findings indicate similar patterns of scene-specific neural reinstatement between children and young adults. Building on Fandakova et al. (2019), who found similar distinctiveness of neural representations during encoding in 8-to-15-year-old children and adults, our results suggest that this similarity extends to younger ages, showing comparable distinctiveness of neural representations for unique memories from middle to late childhood and early adolescence. Additionally, our research supports the presence of scene-specific reinstatement in 5-to-7-year-old children, albeit at a lower level compared to adults, aligning with previous studies (Benear et al., 2022; Cohen et al., 2019; Golarai et al., 2015), which demonstrated reliable mnemonic reinstatement for visual input (i.e., faces, movie clips) in 5-to-11-year-old children. Furthermore, we extend these findings, by showing that successful of long-term memory retrieval is associated with more differentiated neural reinstatement in both children and young adults, indicating similar mechanisms of detail-rich memory consolidation present as early as 5-to-7 year.

Removed in the revised version: In terms of category-based reinstatement, our results showed that only children demonstrated category-based reinstatement of to-be-retrieved memories in medial-temporal, prefrontal, parietal, and cerebellar brain regions, while no category-based reinstatement was observed in occipital brain regions. Furthermore, category-based reinstatement increased from immediate or short delay to long delay in the HC and PHG only in children. Parietal and medial temporal brain regions were reported to carry event-specific information for episodic memories (Jonker et al., 2018; Sekeres et al., 2016; Sestieri et al., 2017; Winocur et al., 2007). However, in children in long delay these regions show neural reinstatement for only gist-like memories, possibly reflecting a faster decay of information specificity that these regions otherwise carry in young adults. As for prefrontal brain regions, gist-like neural reinstatement in this region in children may reflect integration of memory representations into more abstract categorical representations, considering mPFC’s role in integrating across memories (Schlichting et al., 2015), and integrating new information into schema (Gilboa & Marlatte, 2017; Preston & Eichenbaum, 2013).

Removed in the revised version: Overall, these results are in line with Brainerd et al. (2002), showing that in middle childhood exact memories (i.e., scene-specific reinstatement) and gist-memories (i.e., category-based reinstatement) co-exist. In long delay, on the other hand, children’s retrieval was based on less-specific gist-like mnemonic reinstatement in the frontal, cerebellar, medial temporal, and parietal brain regions, in the absence of detail-rich scene-specific memories in these brain regions. This is in line with Keresztes et al. (2018) and Ngo et al. (2021) and extend their postulations, showing that 5-to-7-year-old children tend to rely more on generalization not only during encoding but also during long-delay memory consolidation, which is reflected through category-based neural pattern reinstatement during memory retention over time.

Earlier version: Exploring reinstatement and behavioural performance relationships, our results showed that over time higher scene-specific reinstatement is related to better mnemonic recollection both in children and young adults. This is in line with the neural fidelity hypothesis (Xue, 2018), suggesting that more similar neural reinstatement reflect less noisy representations of mnemonic information. Convergent evidence showed that higher fidelity of neural representation across study episodes leads to successful memory (Xue et al., 2010, 2013). In addition, we found that higher long-delay gist-like reinstatement was related to worse memory retention rates only in children. In general, only children exhibited gist-like reinstatement of scenes, which further increased as memories aged. Convergent with the observed reverse effect of scene-specific and category-based reinstatement on memory recollection, Masís-Obando et al. (2022) reported similar effects in adults. Namely, the authors showed that more specific neural representations predicted subsequent memory performance in young adults. On the other hand, more schema-based representations in posterior HC were related to pooper subsequent performance, while more schema-based representations in mPFC were related to better subsequent memory. However, Masís-Obando et al. (2022) used stimuli with clearly differentiable schema and details components, which may have rendered significant neural representations in young adults. Future studies may use this approach to examine the benefits of schema-based representations, and the age differences therein.

Revised version: Our results indicate that higher scene-specific neural reinstatement over time correlated with better memory retention in both children and adults. This is in line with the neural fidelity hypothesis (Xue, 2018), suggesting that more similar neural reinstatement reflect less noisy representations of mnemonic information. Convergent evidence showed that higher fidelity of neural representation across study episodes leads to successful memory (Xue et al., 2010, 2013). Similarly, Masís-Obando et al. (2022) reported that more specific neural representations predicted subsequent memory performance in young adults.

Added in the revised version: Of note, our study design, which resulted in temporal autocorrelation in the BOLD signal between memory retrieval (i.e., fixation time window) and scene observation and response (i.e., scene time window), was consistent across all three delay windows. Since the retrieval procedure remained unchanged over time and was similarly influenced by temporal autocorrelations, we attribute our RSA findings to differences in reinstatement between recent and remote trials. Given that the scene time window for the 3AFC task was constant, the brain signals should exhibit similar perception-based but variably memory-based patterns across all delays.

Added in the revised version: Furthermore, all items, regardless of retrieval delay, underwent extensive learning and showed successful consolidation, as evidenced by correct recall. This suggests that both the fixation and scene time windows engaged memory-related neural processes. According to Brodt et al., (2016, 2018), rapid consolidation-related neural reorganization can occur immediately after learning, indicating that even during recent retrieval, scenes are processed in a memory-oriented manner. Additionally, during the scene time window, participants engaged in retrieval by selecting the correct object location within the scene. Thus, while the scene time window involved perceptual processing, its impact is consistent across all items due to uniform exposure to repeated learning, making them equally familiar to participants. Although our paradigm per se cannot arbitrate between perception-based and memory-based nature of retrieval during scene presentation, our exploratory univariate analysis during the scene presenations time window (see Figure S3, Table S9 in Supplementary Materials) revealed higher neural engagement in the key memory regions with passing time, supporting memory-related processing during the scene time window.

Earlier version: Taken together, our findings provided novel evidence that children showed more attenuated scene-specific reinstatement compared to young adults. Nevertheless, scene-specific reinstatement was beneficial for memory retention in both groups. At the same time, children showed more category-based reinstatement, which was negatively related to memory retention. Therefore, both reduced differentiated reinstatement and enhanced reliance on gist information are characteristic of children’s memory retrieval across time.

Revised version: Taken together, our findings provide novel evidence that although children exhibit more attenuated scene-specific reinstatement compared to young adults, the consolidation-related decrease in differentiated reinstatement follows similar patterns as in adults. This highlights that despite less robust memory consolidation and lower memory strength, children’s neural transformations of distinct memories over time may share the same mechanisms as adults, with scene-specific reinstatement proving beneficial for memory retention in both groups.

Added in the revised version: Unique Gist-like Reinstatement in Children

Added in the revised version: In terms of more generic gist-like reinstatement, our results showed that only children demonstrated such reinstatement in anterior hippocampal, prefrontal, and parietal brain regions during successful retrieval. Furthermore, higher short-delay gist-like reinstatement in the anterior hippocampus and mPFC was associated with poorer short-delay memory accuracy in children. Similarly, higher long-delay gist-like reinstatement was associated with poorer long-delay memory accuracy in children. With these findings, we provide the first neural empirical evidence to support the Fuzzy Trace Theory (Brainerd & Reyna, 2002; Reyna & Brainerd, 1995), showing neural reorganization of memory representations in children.

Added in the revised version: The Fuzzy Trace Theory aims to characterized the shifts in ongoing balance between precise, detailed “verbatim” memory and more generalize, simplified “gist” memory (Brainerd & Reyna, 2002) from a developmental perspective. Our associative object-location task allowed the investigation of these “dichotomy” as it was aimed to cultivate detailed, precise memories for retrieval. Simultaneously, it enabled generalization by creating of more generic representations due to the presence of related category-based information. Adults were able to build upon solid pre-existing knowledge by embellishing them with details and integrating them into these structures. Children, in contrast, with their sparser knowledge, may have focused more on solidifying the structure with overlapping information. Aligning with the Fuzzy Trace Theory, our results suggest that reliance on gist-like memory representations is less effective for long-term retention of complex associative information compared to detailed verbatim memory, which seems to be characteristic of adults.

Added in the revised version: The association between short-delay, gist-like reinstatement in the anterior hippocampus and mPFC in children align with the findings that, in middle childhood, the anterior hippocampus is generally functionally connected with frontal brain regions and associated with semantic memory (Plachti et al., 2023). Earlier maturation of anterior hippocampus in middle childhood (Canada et al., 2021), along with its more pronounced role in associative memory (Lee et al., 2020), contribute to our understanding of its role in consolidation-related neural reorganization in children. On the other hand, studies with adult subjects show that gist-like reinstatement in posterior hippocampal is linked to more generic semantic gist of the original memory in adults (Dandolo & Schwabe, 2018; Krenz et al., 2023). In line with this, more schema-based representations in posterior hippocampus were related to pooper subsequent performance in adults (Masís-Obando et al., 2022). A more prolonged maturation of the posterior hippocampus, along with the functional shift within the anterior-posterior hippocampal axis with respect to episodic memory, suggest that neural transformations of mnemonic representations in children may be governed by inherently different neural mechanisms. The mechanisms may result in weaker memories for detailed, complex information over long time in children (Canada et al., 2021; Ghetti & Lee, 2013; Plachti et al., 2023). Particularly, pronounced functional connectivity between the anterior hippocampus and frontal regions in children, coupled with less differentiated functional specification and broad cognitive covariance network within these regions (Plachti et al., 2023), may underlie more sparse retention patterns and less differentiated memory reorganization in children.

Added in the revised version: The gist-like neural reinstatement in the mPFC in children may reflect consolidation-related integration of memory representations into more abstract, generic forms. This aligns with the mPFC’s known role in integrating across memories (Schlichting et al., 2015), the increase in semantically transformed representations for related information over time in adults (Krenz et al., 2023), and the integration new information into schema (Gilboa & Marlatte, 2017; Preston & Eichenbaum, 2013). While gist-like neural representations may support the generalization of information to bolster the sparse knowledge structures in children, this occurs at the costs of memory precision (Reyna et al., 2016). Consequently, there is a negative association between gist-like memory reinstatement in the mPFC and memory accuracy both in short and long delay. In contrast to our findings, Masís-Obando et al. (2022) demonstrated that more schema-based representations in the mPFC were associated with better subsequent memory performance in adults. However, the study utilized stimuli with clearly differentiable schema and details components. Future studies may use this approach to further explore these differences and the specific conditions under which schema-based representations enhance memory performance, and the age differences therein.

Added in the revised version: Overall, our results are in line with Brainerd et al. (2002), showing that in middle childhood, precise mnemonic representations (i.e., scene-specific reinstatement) and gist mnemonic representations (i.e., gist-like reinstatement) co-exist also on the neural level. We performed exploratory analysis that revealed a negative relationship between detailed scene-specific reinstatement and generic gist-like reinstatement in children (see section S3.1, Figure S5 in Supplementary Materials). Therefore, children with lower item-specific menminic representations tend to show more generic gist-like representations. Extending on the postulations from Keresztes et al. (2018) and Ngo et al. (2021) that 5-to-7-year-old children tend to rely more on generalization, our findings suggest that retaining memories with less specific details may allow for faster integration of overlapping features into emerging knowledge structures (Bauer, 2021; Gilboa & Marlatte, 2017).

Added in the revised version: On the other hand, adults could form strong, highly detailed memories aided by effective strategic retrieval methods, without the need to form gist-like representations. Moreover, they may have employed different retrieval neural mechanisms than children, as indicated by our exploratory findings that higher neural engagement over time was associated with the decrease in scene-specific neural reinstatement in adults (see section S3.1, Figure S5 in Supplementary Materials); suggesting a higher recruitment of neural resources to compensate for decaying memory reinstatement.

Added in the revised version: Taken together, our findings provide novel evidence that an enhanced reliance on gist information characterizes children’s memory retrieval across time. With this we provide the first empirical evidence to support the Fuzzy trace theory on the level of gist-like neural representations evolvement in children.

Limitations

Earlier version: Several limitations of the current study should be noted. First, our test for memory was based on a 3-alternative forced choice procedure, which was intended to reduce the need for strategic search (e.g., in free recall). As reorganization and stabilization in consolidation depend on the psychological nature of mnemonic representations (Moscovitch & Gilboa, 2022), future studies may employ more demanding recall-based memories to characterize memory consolidation more comprehensively. Particularly, future studies may differentiate mnemonic accessibility vs. precision (Murray et al., 2015; Richter et al., 2016), as they may show differential temporal dynamics in the developing brain and involve differential neural mechanisms. Second, as we included only stimuli congruent with prior knowledge, future studies may introduce knowledge-incongruent information to investigate the beneficial effect of prior knowledge on memory consolidation more directly. Prior knowledge may impact learning and consolidation of information over time differentially by development (McKenzie & Eichenbaum, 2011; van Kesteren et al., 2013; Wang& Morris, 2010). Third, we concentrated on a limited age range in middle childhood. To characterize how neural mechanisms of memory consolidation evolve over time, future studies should include other developmental cohorts. Fourth, although we focused on ROIs associated with the recollection network and implicated in retrieval of visual information, we did not investigate the connectivity between these brain regions and how it changes as memories age. Future studies should investigate consolidation-related neural connectivity patterns and their temporal dynamics in the developing brain. Finally, children in our sample were positively biased in socio-demographical score and IQ compared to young adults, which may restrict the generalizability of our results.

Revised version: Several limitations of the current study should be noted. First, our test for memory was based on a 3-alternative forced choice procedure, which was intended to reduce the need for strategic search (e.g., in free recall). As reorganization and stabilization in consolidation depend on the psychological nature of mnemonic representations (Moscovitch & Gilboa, 2022), future studies may employ more demanding recall-based memories to characterize memory consolidation more comprehensively. Particularly, future studies may differentiate mnemonic accessibility vs. precision (Murray et al., 2015; Richter et al., 2016), as they may show differential temporal dynamics in the developing brain and involve differential neural mechanisms. Second, as we included only stimuli congruent with prior knowledge, future studies may introduce knowledge-incongruent information to investigate the beneficial effect of prior knowledge on memory consolidation more directly. Prior knowledge may impact learning and consolidation of information over time differentially by development (McKenzie & Eichenbaum, 2011; van Kesteren et al., 2013; Wang& Morris, 2010). Third, we concentrated on a limited age range in middle childhood. To characterize how neural mechanisms of memory consolidation evolve over time, future studies should include other developmental cohorts. Fourth, we acknowledge that our study design leads to temporal autocorrelation in the BOLD signal when calculating RSA between fixation and scene time windows. Although we argue that our results, given the identical procedure over time, are more attributed to the delay-related changes in the neural reinstatement, future studies should tailor the design of the retrieval procedure to warrant cross-run comparisons. This could be achieved by introducing the same items repeatedly across different runs. Fifth, our task may not have been demanding enough for young adults to fully challenge their memory retention and encourage the formation of more gist-like representations. Future studies could explore this further by using more challenging conditions to enhance the formation of more generic memories in adults, avoid bias related to prior knowledge. Sixth, although we focused on ROIs associated with the recollection network and implicated in retrieval of visual information, we did not investigate the connectivity between these brain regions and how it changes as memories age. Future studies should investigate consolidation-related neural connectivity patterns and their temporal dynamics in the developing brain. Finally, children in our sample were positively biased in socio-demographical score and IQ compared to young adults, which may restrict the generalizability of our results.

Conclusions

Earlier version: In this study, we present novel empirical evidence regarding the neural mechanisms underlying less robust memory retention of intentionally learned object-location associations in 5-to-7-year-old children compared to young adults. Our findings indicate that over time, children exhibited less activation in core recollection and cerebellar brain regions during successful retrieval of remote memories. Additionally, they showed less differentiated neural reinstatement over time, but enhanced category-based reinstatement of retrieved memories. Taken together, our results suggest that compared to younger adults’ mature consolidation system, memories of the developing brain in early school years are characterized by attenuated specificity of mnemonic representations and lower activation during retrieval in core retrieval brain network regions. At the same time, gist-based, schema representations play an important role in children’s retrieval. This knowledge could potentially be used to design interventions that focus on enhancing long-term memory retention in children. Specifically, memory training programs may be tailored towards young school children to promote strategies that encode and retrieve specific and detailed representations of memories by capitalizing and making explicit connections to relevant gist-based, schema knowledge.

Revised version: In this study, we present novel empirical evidence on the neural mechanisms underlying the less robust memory retention of intentionally learned object-location associations in 5-to-7-year-old children compared to young adults. Our findings reveal that over time, children show attenuated consolidation-related upregulation in neocortical and cerebellar brain regions during successful retrieval. Furthermore, children may form different neural memory representations than young adults, as evidenced by the coexistence of both detailed scene-specific and generic gist-like reinstatement. Our results suggest that, unlike the mature consolidation systems in young adults, the developing brains of early school-age children show attenuated yet adult-like item-specific representations and reduced neural upregulation in core retrieval networks. Additionally, gist-based representations play a significant role in children’s retrieval processes, possibly aiding the building up of schema knowledge.

Data availability statement

The datasets generated and analysed during the current study are available from the corresponding authors upon reasonable request.

Conflict of interest disclosure

We have no known conflict of interest to disclose.

Supporting information